{
  "./goals/whitebox/whiteBoxTruthTest0422.adimen.tstp": [
    "ProverResult.PROOF_FOUND",
    "b\"# No SInE strategy applied\\n# Auto-Mode selected heuristic SAT001_MinMin_NN\\n# and selection function SelectMaxLComplexAvoidPosPred.\\n#\\n# Presaturation interreduction done\\n\\n# Proof found!\\n# SZS status Theorem\\n# SZS output start CNFRefutation\\nfof(schemaBinaryRelationA18, axiom, ![X1033, X1034]:(p__d__holds3(c__equal,X1033,X1034)<=>X1033=X1034), file('adimen.sumo.tstp', schemaBinaryRelationA18)).\\nfof(whiteBoxTruthTest0422, conjecture, ?[X72]:(p__d__instance(X72,c__BinaryRelation)&![X18, X19]:(~(p__d__holds3(X72,X18,X19))|p__d__holds3(X72,X19,X18))), file('./goals/whitebox/whiteBoxTruthTest0422.adimen.tstp', whiteBoxTruthTest0422)).\\nfof(typeA3, axiom, ![X1038, X1039, X1040]:(p__d__holds3(X1038,X1039,X1040)=>p__d__instance(X1038,c__BinaryRelation)), file('adimen.sumo.tstp', typeA3)).\\nfof(c_0_3, plain, ![X12112, X12113]:((~p__d__holds3(c__equal,X12112,X12113)|X12112=X12113)&(X12112!=X12113|p__d__holds3(c__equal,X12112,X12113))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[schemaBinaryRelationA18])])).\\nfof(c_0_4, negated_conjecture, ~(?[X72]:(p__d__instance(X72,c__BinaryRelation)&![X18, X19]:(~p__d__holds3(X72,X18,X19)|p__d__holds3(X72,X19,X18)))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[whiteBoxTruthTest0422])])).\\nfof(c_0_5, plain, ![X13340, X13341, X13342]:(~p__d__holds3(X13340,X13341,X13342)|p__d__instance(X13340,c__BinaryRelation)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[typeA3])])).\\ncnf(c_0_6, plain, (p__d__holds3(c__equal,X1,X2)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_3])).\\nfof(c_0_7, negated_conjecture, ![X14609]:((p__d__holds3(X14609,esk1216_1(X14609),esk1217_1(X14609))|~p__d__instance(X14609,c__BinaryRelation))&(~p__d__holds3(X14609,esk1217_1(X14609),esk1216_1(X14609))|~p__d__instance(X14609,c__BinaryRelation))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).\\ncnf(c_0_8, plain, (p__d__instance(X1,c__BinaryRelation)|~p__d__holds3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).\\ncnf(c_0_9, plain, (p__d__holds3(c__equal,X1,X1)), inference(er,[status(thm)],[c_0_6])).\\ncnf(c_0_10, plain, (X1=X2|~p__d__holds3(c__equal,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_3])).\\ncnf(c_0_11, negated_conjecture, (p__d__holds3(X1,esk1216_1(X1),esk1217_1(X1))|~p__d__instance(X1,c__BinaryRelation)), inference(split_conjunct,[status(thm)],[c_0_7])).\\ncnf(c_0_12, plain, (p__d__instance(c__equal,c__BinaryRelation)), inference(spm,[status(thm)],[c_0_8, c_0_9])).\\ncnf(c_0_13, negated_conjecture, (~p__d__holds3(X1,esk1217_1(X1),esk1216_1(X1))|~p__d__instance(X1,c__BinaryRelation)), inference(split_conjunct,[status(thm)],[c_0_7])).\\ncnf(c_0_14, negated_conjecture, (esk1217_1(c__equal)=esk1216_1(c__equal)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).\\ncnf(c_0_15, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_9]), c_0_12])]), ['proof']).\\n# SZS output end CNFRefutation\\n# Training examples: 0 positive, 0 negative\\n\""
  ],
  "./goals/whitebox/whiteBoxTruthTest7835.adimen.tstp": [
    "ProverResult.PROOF_FOUND",
    "b\"# No SInE strategy applied\\n# Auto-Mode selected heuristic SAT001_MinMin_NN\\n# and selection function SelectMaxLComplexAvoidPosPred.\\n#\\n# Presaturation interreduction done\\n\\n# Proof found!\\n# SZS status Theorem\\n# SZS output start CNFRefutation\\nfof(whiteBoxTruthTest7835, conjecture, ?[X47]:~(p__d__instance(X47,c__Francium)), file('./goals/whitebox/whiteBoxTruthTest7835.adimen.tstp', whiteBoxTruthTest7835)).\\nfof(predefinitionsA12, axiom, ![X1, X2, X3]:((p__d__instance(X1,X2)&p__d__subclass(X2,X3))=>p__d__instance(X1,X3)), file('adimen.sumo.tstp', predefinitionsA12)).\\nfof(miloA4420, axiom, p__d__subclass(c__Francium,c__ElementalSubstance), file('adimen.sumo.tstp', miloA4420)).\\nfof(predefinitionsA15, axiom, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~(p__d__instance(X6,X4))|~(p__d__instance(X6,X5)))), file('adimen.sumo.tstp', predefinitionsA15)).\\nfof(mergeA235, axiom, p__d__subclass(c__ElementalSubstance,c__PureSubstance), file('adimen.sumo.tstp', mergeA235)).\\nfof(mergeA233, axiom, p__d__subclass(c__PureSubstance,c__Substance), file('adimen.sumo.tstp', mergeA233)).\\nfof(mergeA257, axiom, p__d__disjoint(c__Collection,c__SelfConnectedObject), file('adimen.sumo.tstp', mergeA257)).\\nfof(mergeA226, axiom, p__d__subclass(c__Substance,c__SelfConnectedObject), file('adimen.sumo.tstp', mergeA226)).\\nfof(mergeA3229, axiom, p__d__subclass(c__Group,c__Collection), file('adimen.sumo.tstp', mergeA3229)).\\nfof(miloA4176, axiom, ![X1]:(p__d__subclass(X1,c__UniformClothing)=>?[X901]:(p__d__instance(X901,c__Group)&p__hasUniform(X901,X1))), file('adimen.sumo.tstp', miloA4176)).\\nfof(predefinitionsA7, axiom, ![X1]:p__d__subclass(X1,X1), file('adimen.sumo.tstp', predefinitionsA7)).\\nfof(c_0_11, negated_conjecture, ~(?[X47]:~p__d__instance(X47,c__Francium)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[whiteBoxTruthTest7835])])).\\nfof(c_0_12, plain, ![X7776, X7777, X7778]:(~p__d__instance(X7776,X7777)|~p__d__subclass(X7777,X7778)|p__d__instance(X7776,X7778)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA12])])).\\nfof(c_0_13, negated_conjecture, ![X14607]:p__d__instance(X14607,c__Francium), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])).\\ncnf(c_0_14, plain, (p__d__instance(X1,X3)|~p__d__instance(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).\\ncnf(c_0_15, plain, (p__d__subclass(c__Francium,c__ElementalSubstance)), inference(split_conjunct,[status(thm)],[miloA4420])).\\ncnf(c_0_16, negated_conjecture, (p__d__instance(X1,c__Francium)), inference(split_conjunct,[status(thm)],[c_0_13])).\\nfof(c_0_17, plain, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~p__d__instance(X6,X4)|~p__d__instance(X6,X5))), inference(fof_simplification,[status(thm)],[predefinitionsA15])).\\ncnf(c_0_18, plain, (p__d__subclass(c__ElementalSubstance,c__PureSubstance)), inference(split_conjunct,[status(thm)],[mergeA235])).\\ncnf(c_0_19, plain, (p__d__instance(X1,c__ElementalSubstance)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])])).\\nfof(c_0_20, plain, ![X7779, X7780, X7781, X7782, X7783]:((~p__d__disjoint(X7779,X7780)|(~p__d__instance(X7781,X7779)|~p__d__instance(X7781,X7780)))&((p__d__instance(esk1_2(X7782,X7783),X7782)|p__d__disjoint(X7782,X7783))&(p__d__instance(esk1_2(X7782,X7783),X7783)|p__d__disjoint(X7782,X7783)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])])])).\\ncnf(c_0_21, plain, (p__d__subclass(c__PureSubstance,c__Substance)), inference(split_conjunct,[status(thm)],[mergeA233])).\\ncnf(c_0_22, plain, (p__d__instance(X1,c__PureSubstance)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_18]), c_0_19])])).\\ncnf(c_0_23, plain, (~p__d__disjoint(X1,X2)|~p__d__instance(X3,X1)|~p__d__instance(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).\\ncnf(c_0_24, plain, (p__d__disjoint(c__Collection,c__SelfConnectedObject)), inference(split_conjunct,[status(thm)],[mergeA257])).\\ncnf(c_0_25, plain, (p__d__subclass(c__Substance,c__SelfConnectedObject)), inference(split_conjunct,[status(thm)],[mergeA226])).\\ncnf(c_0_26, plain, (p__d__instance(X1,c__Substance)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_21]), c_0_22])])).\\ncnf(c_0_27, plain, (~p__d__instance(X1,c__SelfConnectedObject)|~p__d__instance(X1,c__Collection)), inference(spm,[status(thm)],[c_0_23, c_0_24])).\\ncnf(c_0_28, plain, (p__d__instance(X1,c__SelfConnectedObject)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_25]), c_0_26])])).\\ncnf(c_0_29, plain, (p__d__subclass(c__Group,c__Collection)), inference(split_conjunct,[status(thm)],[mergeA3229])).\\ncnf(c_0_30, plain, (~p__d__instance(X1,c__Collection)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28])])).\\nfof(c_0_31, plain, ![X11978]:((p__d__instance(esk1193_1(X11978),c__Group)|~p__d__subclass(X11978,c__UniformClothing))&(p__hasUniform(esk1193_1(X11978),X11978)|~p__d__subclass(X11978,c__UniformClothing))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[miloA4176])])])])).\\ncnf(c_0_32, plain, (~p__d__instance(X1,c__Group)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_29]), c_0_30])).\\ncnf(c_0_33, plain, (p__d__instance(esk1193_1(X1),c__Group)|~p__d__subclass(X1,c__UniformClothing)), inference(split_conjunct,[status(thm)],[c_0_31])).\\nfof(c_0_34, plain, ![X7770]:p__d__subclass(X7770,X7770), inference(variable_rename,[status(thm)],[predefinitionsA7])).\\ncnf(c_0_35, plain, (~p__d__subclass(X1,c__UniformClothing)), inference(spm,[status(thm)],[c_0_32, c_0_33])).\\ncnf(c_0_36, plain, (p__d__subclass(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).\\ncnf(c_0_37, plain, ($false), inference(spm,[status(thm)],[c_0_35, c_0_36]), ['proof']).\\n# SZS output end CNFRefutation\\n# Training examples: 0 positive, 0 negative\\n\""
  ],
  "./goals/whitebox/whiteBoxTruthTest4270.adimen.tstp": [
    "ProverResult.PROOF_FOUND",
    "b\"# No SInE strategy applied\\n# Auto-Mode selected heuristic SAT001_MinMin_NN\\n# and selection function SelectMaxLComplexAvoidPosPred.\\n#\\n# Presaturation interreduction done\\n# SatCheck found unsatisfiable ground set\\n\\n# Proof found!\\n# SZS status Theorem\\n# SZS output start CNFRefutation\\nfof(whiteBoxTruthTest4270, conjecture, ?[X172]:~(p__d__instance(X172,c__TwoDimensionalFigure)), file('./goals/whitebox/whiteBoxTruthTest4270.adimen.tstp', whiteBoxTruthTest4270)).\\nfof(predefinitionsA12, axiom, ![X1, X2, X3]:((p__d__instance(X1,X2)&p__d__subclass(X2,X3))=>p__d__instance(X1,X3)), file('adimen.sumo.tstp', predefinitionsA12)).\\nfof(mergeA3611, axiom, p__d__subclass(c__TwoDimensionalFigure,c__GeometricFigure), file('adimen.sumo.tstp', mergeA3611)).\\nfof(predefinitionsA15, axiom, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~(p__d__instance(X6,X4))|~(p__d__instance(X6,X5)))), file('adimen.sumo.tstp', predefinitionsA15)).\\nfof(predefinitionsA24, axiom, ((((![X7, X8, X9]:(p__d__disjointDecomposition3(X7,X8,X9)<=>p__d__disjoint(X8,X9))&![X7, X8, X9, X10]:(p__d__disjointDecomposition4(X7,X8,X9,X10)<=>((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X9,X10))))&![X7, X8, X9, X10, X11]:(p__d__disjointDecomposition5(X7,X8,X9,X10,X11)<=>(((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X10,X11))))&![X7, X8, X9, X10, X11, X12]:(p__d__disjointDecomposition6(X7,X8,X9,X10,X11,X12)<=>(((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X11,X12))))&![X7, X8, X9, X10, X11, X12, X13]:(p__d__disjointDecomposition7(X7,X8,X9,X10,X11,X12,X13)<=>((((((((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X8,X13))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X9,X13))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X10,X13))&p__d__disjoint(X11,X12))&p__d__disjoint(X11,X13))&p__d__disjoint(X12,X13)))), file('adimen.sumo.tstp', predefinitionsA24)).\\nfof(predefinitionsA8, axiom, ![X1, X2, X3]:((p__d__subclass(X1,X2)&p__d__subclass(X2,X3))=>p__d__subclass(X1,X3)), file('adimen.sumo.tstp', predefinitionsA8)).\\nfof(mergeA3607, axiom, p__d__subclass(c__GeometricFigure,c__ShapeAttribute), file('adimen.sumo.tstp', mergeA3607)).\\nfof(mergeA331, axiom, p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition), file('adimen.sumo.tstp', mergeA331)).\\nfof(mergeA333, axiom, p__d__subclass(c__Attribute,c__Abstract), file('adimen.sumo.tstp', mergeA333)).\\nfof(mergeA330, axiom, p__d__subclass(c__Abstract,c__Entity), file('adimen.sumo.tstp', mergeA330)).\\nfof(mergeA351, axiom, p__d__subclass(c__InternalAttribute,c__Attribute), file('adimen.sumo.tstp', mergeA351)).\\nfof(mergeA3596, axiom, p__d__subclass(c__ShapeAttribute,c__StructureAttribute), file('adimen.sumo.tstp', mergeA3596)).\\nfof(mergeA3595, axiom, p__d__subclass(c__StructureAttribute,c__InternalAttribute), file('adimen.sumo.tstp', mergeA3595)).\\nfof(mergeA575, axiom, p__d__subclass(c__Class,c__SetOrClass), file('adimen.sumo.tstp', mergeA575)).\\nfof(mergeA177, axiom, ![X7]:(p__d__instance(X7,c__Class)<=>p__d__subclass(X7,c__Entity)), file('adimen.sumo.tstp', mergeA177)).\\nfof(c_0_15, negated_conjecture, ~(?[X172]:~p__d__instance(X172,c__TwoDimensionalFigure)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[whiteBoxTruthTest4270])])).\\nfof(c_0_16, plain, ![X7776, X7777, X7778]:(~p__d__instance(X7776,X7777)|~p__d__subclass(X7777,X7778)|p__d__instance(X7776,X7778)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA12])])).\\nfof(c_0_17, negated_conjecture, ![X14607]:p__d__instance(X14607,c__TwoDimensionalFigure), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).\\ncnf(c_0_18, plain, (p__d__instance(X1,X3)|~p__d__instance(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).\\ncnf(c_0_19, plain, (p__d__subclass(c__TwoDimensionalFigure,c__GeometricFigure)), inference(split_conjunct,[status(thm)],[mergeA3611])).\\ncnf(c_0_20, negated_conjecture, (p__d__instance(X1,c__TwoDimensionalFigure)), inference(split_conjunct,[status(thm)],[c_0_17])).\\nfof(c_0_21, plain, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~p__d__instance(X6,X4)|~p__d__instance(X6,X5))), inference(fof_simplification,[status(thm)],[predefinitionsA15])).\\nfof(c_0_22, plain, ![X7870, X7871, X7872, X7873, X7874, X7875, X7876, X7877, X7878, X7879, X7880, X7881, X7882, X7883, X7884, X7885, X7886, X7887, X7888, X7889, X7890, X7891, X7892, X7893, X7894]:((((((~p__d__disjointDecomposition3(X7870,X7871,X7872)|p__d__disjoint(X7871,X7872))&(~p__d__disjoint(X7871,X7872)|p__d__disjointDecomposition3(X7870,X7871,X7872)))&((((p__d__disjoint(X7874,X7875)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))&(p__d__disjoint(X7874,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(p__d__disjoint(X7875,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(~p__d__disjoint(X7874,X7875)|~p__d__disjoint(X7874,X7876)|~p__d__disjoint(X7875,X7876)|p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))))&(((((((p__d__disjoint(X7878,X7879)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))&(p__d__disjoint(X7878,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7878,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7880,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(~p__d__disjoint(X7878,X7879)|~p__d__disjoint(X7878,X7880)|~p__d__disjoint(X7878,X7881)|~p__d__disjoint(X7879,X7880)|~p__d__disjoint(X7879,X7881)|~p__d__disjoint(X7880,X7881)|p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))))&(((((((((((p__d__disjoint(X7883,X7884)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))&(p__d__disjoint(X7883,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7886,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(~p__d__disjoint(X7883,X7884)|~p__d__disjoint(X7883,X7885)|~p__d__disjoint(X7883,X7886)|~p__d__disjoint(X7883,X7887)|~p__d__disjoint(X7884,X7885)|~p__d__disjoint(X7884,X7886)|~p__d__disjoint(X7884,X7887)|~p__d__disjoint(X7885,X7886)|~p__d__disjoint(X7885,X7887)|~p__d__disjoint(X7886,X7887)|p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))))&((((((((((((((((p__d__disjoint(X7889,X7890)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894))&(p__d__disjoint(X7889,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7893,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(~p__d__disjoint(X7889,X7890)|~p__d__disjoint(X7889,X7891)|~p__d__disjoint(X7889,X7892)|~p__d__disjoint(X7889,X7893)|~p__d__disjoint(X7889,X7894)|~p__d__disjoint(X7890,X7891)|~p__d__disjoint(X7890,X7892)|~p__d__disjoint(X7890,X7893)|~p__d__disjoint(X7890,X7894)|~p__d__disjoint(X7891,X7892)|~p__d__disjoint(X7891,X7893)|~p__d__disjoint(X7891,X7894)|~p__d__disjoint(X7892,X7893)|~p__d__disjoint(X7892,X7894)|~p__d__disjoint(X7893,X7894)|p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA24])])])])).\\nfof(c_0_23, plain, ![X7771, X7772, X7773]:(~p__d__subclass(X7771,X7772)|~p__d__subclass(X7772,X7773)|p__d__subclass(X7771,X7773)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA8])])).\\ncnf(c_0_24, plain, (p__d__subclass(c__GeometricFigure,c__ShapeAttribute)), inference(split_conjunct,[status(thm)],[mergeA3607])).\\ncnf(c_0_25, plain, (p__d__instance(X1,c__GeometricFigure)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])])).\\nfof(c_0_26, plain, ![X7779, X7780, X7781, X7782, X7783]:((~p__d__disjoint(X7779,X7780)|(~p__d__instance(X7781,X7779)|~p__d__instance(X7781,X7780)))&((p__d__instance(esk1_2(X7782,X7783),X7782)|p__d__disjoint(X7782,X7783))&(p__d__instance(esk1_2(X7782,X7783),X7783)|p__d__disjoint(X7782,X7783)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])])).\\ncnf(c_0_27, plain, (p__d__disjoint(X1,X2)|~p__d__disjointDecomposition6(X3,X4,X1,X2,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_22])).\\ncnf(c_0_28, plain, (p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition)), inference(split_conjunct,[status(thm)],[mergeA331])).\\ncnf(c_0_29, plain, (p__d__subclass(X1,X3)|~p__d__subclass(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).\\ncnf(c_0_30, plain, (p__d__subclass(c__Attribute,c__Abstract)), inference(split_conjunct,[status(thm)],[mergeA333])).\\ncnf(c_0_31, plain, (p__d__subclass(c__Abstract,c__Entity)), inference(split_conjunct,[status(thm)],[mergeA330])).\\ncnf(c_0_32, plain, (p__d__subclass(c__InternalAttribute,c__Attribute)), inference(split_conjunct,[status(thm)],[mergeA351])).\\ncnf(c_0_33, plain, (p__d__subclass(c__ShapeAttribute,c__StructureAttribute)), inference(split_conjunct,[status(thm)],[mergeA3596])).\\ncnf(c_0_34, plain, (p__d__subclass(c__StructureAttribute,c__InternalAttribute)), inference(split_conjunct,[status(thm)],[mergeA3595])).\\ncnf(c_0_35, plain, (p__d__subclass(c__Class,c__SetOrClass)), inference(split_conjunct,[status(thm)],[mergeA575])).\\ncnf(c_0_36, plain, (p__d__instance(X1,c__ShapeAttribute)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_24]), c_0_25])])).\\ncnf(c_0_37, plain, (~p__d__disjoint(X1,X2)|~p__d__instance(X3,X1)|~p__d__instance(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).\\ncnf(c_0_38, plain, (p__d__disjoint(c__Attribute,c__SetOrClass)), inference(spm,[status(thm)],[c_0_27, c_0_28])).\\nfof(c_0_39, plain, ![X8007]:((~p__d__instance(X8007,c__Class)|p__d__subclass(X8007,c__Entity))&(~p__d__subclass(X8007,c__Entity)|p__d__instance(X8007,c__Class))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mergeA177])])).\\ncnf(c_0_40, plain, (p__d__subclass(X1,c__Abstract)|~p__d__subclass(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_29, c_0_30])).\\ncnf(c_0_41, plain, (p__d__subclass(X1,c__Entity)|~p__d__subclass(X1,c__Abstract)), inference(spm,[status(thm)],[c_0_29, c_0_31])).\\ncnf(c_0_42, plain, (p__d__subclass(X1,c__Attribute)|~p__d__subclass(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_29, c_0_32])).\\ncnf(c_0_43, plain, (p__d__subclass(X1,c__StructureAttribute)|~p__d__subclass(X1,c__ShapeAttribute)), inference(spm,[status(thm)],[c_0_29, c_0_33])).\\ncnf(c_0_44, plain, (p__d__subclass(X1,c__InternalAttribute)|~p__d__subclass(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_29, c_0_34])).\\ncnf(c_0_45, plain, (p__d__subclass(X1,c__ShapeAttribute)|~p__d__subclass(X1,c__GeometricFigure)), inference(spm,[status(thm)],[c_0_29, c_0_24])).\\ncnf(c_0_46, plain, (p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Class)), inference(spm,[status(thm)],[c_0_18, c_0_35])).\\ncnf(c_0_47, plain, (p__d__instance(X1,c__Attribute)|~p__d__instance(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_18, c_0_32])).\\ncnf(c_0_48, plain, (p__d__instance(X1,c__StructureAttribute)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_33]), c_0_36])])).\\ncnf(c_0_49, plain, (p__d__instance(X1,c__InternalAttribute)|~p__d__instance(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_18, c_0_34])).\\ncnf(c_0_50, plain, (~p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_37, c_0_38])).\\ncnf(c_0_51, plain, (p__d__instance(X1,c__Class)|~p__d__subclass(X1,c__Entity)), inference(split_conjunct,[status(thm)],[c_0_39])).\\ncnf(c_0_52, plain, ($false), inference(cdclpropres,[status(thm)],[c_0_40, c_0_41, c_0_42, c_0_43, c_0_44, c_0_45, c_0_46, c_0_47, c_0_48, c_0_49, c_0_50, c_0_51, c_0_19]), ['proof']).\\n# SZS output end CNFRefutation\\n# Training examples: 0 positive, 0 negative\\n\""
  ],
  "./goals/whitebox/whiteBoxTruthTest6595.adimen.tstp": [
    "ProverResult.PROOF_FOUND",
    "b\"# No SInE strategy applied\\n# Auto-Mode selected heuristic SAT001_MinMin_NN\\n# and selection function SelectMaxLComplexAvoidPosPred.\\n#\\n# Presaturation interreduction done\\n# SatCheck found unsatisfiable ground set\\n\\n# Proof found!\\n# SZS status Theorem\\n# SZS output start CNFRefutation\\nfof(whiteBoxTruthTest6595, conjecture, ?[X208, X908, X909]:(((p__d__instance(X908,c__NightTime)&p__d__instance(X909,c__NightTime))&p__meetsTemporally(X208,X908))&p__meetsTemporally(X909,X208)), file('./goals/whitebox/whiteBoxTruthTest6595.adimen.tstp', whiteBoxTruthTest6595)).\\nfof(miloA2753, axiom, ![X208]:(p__d__instance(X208,c__DayTime)=>?[X908, X909]:(((p__d__instance(X908,c__NightTime)&p__d__instance(X909,c__NightTime))&p__meetsTemporally(X208,X908))&p__meetsTemporally(X909,X208))), file('adimen.sumo.tstp', miloA2753)).\\nfof(miloA2750, axiom, ![X906]:(p__d__instance(X906,c__NightTime)=>?[X200, X201]:(((p__d__instance(X200,c__DayTime)&p__d__instance(X201,c__DayTime))&p__meetsTemporally(X906,X200))&p__meetsTemporally(X201,X906))), file('adimen.sumo.tstp', miloA2750)).\\nfof(predefinitionsA8, axiom, ![X1, X2, X3]:((p__d__subclass(X1,X2)&p__d__subclass(X2,X3))=>p__d__subclass(X1,X3)), file('adimen.sumo.tstp', predefinitionsA8)).\\nfof(mergeA176, axiom, ![X7]:(p__d__subclass(X7,c__Entity)=>?[X35]:p__d__instance(X35,X7)), file('adimen.sumo.tstp', mergeA176)).\\nfof(mergeA330, axiom, p__d__subclass(c__Abstract,c__Entity), file('adimen.sumo.tstp', mergeA330)).\\nfof(mergeA566, axiom, p__d__subclass(c__TimePosition,c__TimeMeasure), file('adimen.sumo.tstp', mergeA566)).\\nfof(mergeA332, axiom, p__d__subclass(c__Quantity,c__Abstract), file('adimen.sumo.tstp', mergeA332)).\\nfof(mergeA561, axiom, p__d__subclass(c__PhysicalQuantity,c__Quantity), file('adimen.sumo.tstp', mergeA561)).\\nfof(mergeA563, axiom, p__d__subclass(c__ConstantQuantity,c__PhysicalQuantity), file('adimen.sumo.tstp', mergeA563)).\\nfof(mergeA564, axiom, p__d__subclass(c__TimeMeasure,c__ConstantQuantity), file('adimen.sumo.tstp', mergeA564)).\\nfof(mergeA568, axiom, p__d__subclass(c__TimeInterval,c__TimePosition), file('adimen.sumo.tstp', mergeA568)).\\nfof(miloA2749, axiom, p__d__subclass(c__NightTime,c__TimeInterval), file('adimen.sumo.tstp', miloA2749)).\\nfof(c_0_13, negated_conjecture, ~(?[X208, X908, X909]:(((p__d__instance(X908,c__NightTime)&p__d__instance(X909,c__NightTime))&p__meetsTemporally(X208,X908))&p__meetsTemporally(X909,X208))), inference(assume_negation,[status(cth)],[whiteBoxTruthTest6595])).\\nfof(c_0_14, negated_conjecture, ![X14609, X14610, X14611]:(~p__d__instance(X14610,c__NightTime)|~p__d__instance(X14611,c__NightTime)|~p__meetsTemporally(X14609,X14610)|~p__meetsTemporally(X14611,X14609)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])).\\nfof(c_0_15, plain, ![X11332]:((((p__d__instance(esk999_1(X11332),c__NightTime)|~p__d__instance(X11332,c__DayTime))&(p__d__instance(esk1000_1(X11332),c__NightTime)|~p__d__instance(X11332,c__DayTime)))&(p__meetsTemporally(X11332,esk999_1(X11332))|~p__d__instance(X11332,c__DayTime)))&(p__meetsTemporally(esk1000_1(X11332),X11332)|~p__d__instance(X11332,c__DayTime))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[miloA2753])])])])).\\ncnf(c_0_16, negated_conjecture, (~p__d__instance(X1,c__NightTime)|~p__d__instance(X2,c__NightTime)|~p__meetsTemporally(X3,X1)|~p__meetsTemporally(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).\\ncnf(c_0_17, plain, (p__meetsTemporally(X1,esk999_1(X1))|~p__d__instance(X1,c__DayTime)), inference(split_conjunct,[status(thm)],[c_0_15])).\\ncnf(c_0_18, plain, (p__d__instance(esk999_1(X1),c__NightTime)|~p__d__instance(X1,c__DayTime)), inference(split_conjunct,[status(thm)],[c_0_15])).\\ncnf(c_0_19, negated_conjecture, (~p__meetsTemporally(X1,X2)|~p__d__instance(X1,c__NightTime)|~p__d__instance(X2,c__DayTime)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])).\\ncnf(c_0_20, plain, (p__meetsTemporally(esk1000_1(X1),X1)|~p__d__instance(X1,c__DayTime)), inference(split_conjunct,[status(thm)],[c_0_15])).\\ncnf(c_0_21, plain, (p__d__instance(esk1000_1(X1),c__NightTime)|~p__d__instance(X1,c__DayTime)), inference(split_conjunct,[status(thm)],[c_0_15])).\\nfof(c_0_22, plain, ![X11326]:((((p__d__instance(esk995_1(X11326),c__DayTime)|~p__d__instance(X11326,c__NightTime))&(p__d__instance(esk996_1(X11326),c__DayTime)|~p__d__instance(X11326,c__NightTime)))&(p__meetsTemporally(X11326,esk995_1(X11326))|~p__d__instance(X11326,c__NightTime)))&(p__meetsTemporally(esk996_1(X11326),X11326)|~p__d__instance(X11326,c__NightTime))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[miloA2750])])])])).\\nfof(c_0_23, plain, ![X7773, X7774, X7775]:(~p__d__subclass(X7773,X7774)|~p__d__subclass(X7774,X7775)|p__d__subclass(X7773,X7775)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA8])])).\\ncnf(c_0_24, negated_conjecture, (~p__d__instance(X1,c__DayTime)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).\\ncnf(c_0_25, plain, (p__d__instance(esk996_1(X1),c__DayTime)|~p__d__instance(X1,c__NightTime)), inference(split_conjunct,[status(thm)],[c_0_22])).\\nfof(c_0_26, plain, ![X8007]:(~p__d__subclass(X8007,c__Entity)|p__d__instance(esk13_1(X8007),X8007)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mergeA176])])])).\\ncnf(c_0_27, plain, (p__d__subclass(X1,X3)|~p__d__subclass(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).\\ncnf(c_0_28, plain, (p__d__subclass(c__Abstract,c__Entity)), inference(split_conjunct,[status(thm)],[mergeA330])).\\ncnf(c_0_29, plain, (p__d__subclass(c__TimePosition,c__TimeMeasure)), inference(split_conjunct,[status(thm)],[mergeA566])).\\ncnf(c_0_30, plain, (p__d__subclass(c__Quantity,c__Abstract)), inference(split_conjunct,[status(thm)],[mergeA332])).\\ncnf(c_0_31, plain, (p__d__subclass(c__PhysicalQuantity,c__Quantity)), inference(split_conjunct,[status(thm)],[mergeA561])).\\ncnf(c_0_32, plain, (p__d__subclass(c__ConstantQuantity,c__PhysicalQuantity)), inference(split_conjunct,[status(thm)],[mergeA563])).\\ncnf(c_0_33, plain, (p__d__subclass(c__TimeMeasure,c__ConstantQuantity)), inference(split_conjunct,[status(thm)],[mergeA564])).\\ncnf(c_0_34, plain, (p__d__subclass(c__TimeInterval,c__TimePosition)), inference(split_conjunct,[status(thm)],[mergeA568])).\\ncnf(c_0_35, negated_conjecture, (~p__d__instance(X1,c__NightTime)), inference(spm,[status(thm)],[c_0_24, c_0_25])).\\ncnf(c_0_36, plain, (p__d__instance(esk13_1(X1),X1)|~p__d__subclass(X1,c__Entity)), inference(split_conjunct,[status(thm)],[c_0_26])).\\ncnf(c_0_37, plain, (p__d__subclass(X1,c__Entity)|~p__d__subclass(X1,c__Abstract)), inference(spm,[status(thm)],[c_0_27, c_0_28])).\\ncnf(c_0_38, plain, (p__d__subclass(X1,c__TimeMeasure)|~p__d__subclass(X1,c__TimePosition)), inference(spm,[status(thm)],[c_0_27, c_0_29])).\\ncnf(c_0_39, plain, (p__d__subclass(X1,c__Abstract)|~p__d__subclass(X1,c__Quantity)), inference(spm,[status(thm)],[c_0_27, c_0_30])).\\ncnf(c_0_40, plain, (p__d__subclass(X1,c__Quantity)|~p__d__subclass(X1,c__PhysicalQuantity)), inference(spm,[status(thm)],[c_0_27, c_0_31])).\\ncnf(c_0_41, plain, (p__d__subclass(X1,c__PhysicalQuantity)|~p__d__subclass(X1,c__ConstantQuantity)), inference(spm,[status(thm)],[c_0_27, c_0_32])).\\ncnf(c_0_42, plain, (p__d__subclass(X1,c__ConstantQuantity)|~p__d__subclass(X1,c__TimeMeasure)), inference(spm,[status(thm)],[c_0_27, c_0_33])).\\ncnf(c_0_43, plain, (p__d__subclass(X1,c__TimePosition)|~p__d__subclass(X1,c__TimeInterval)), inference(spm,[status(thm)],[c_0_27, c_0_34])).\\ncnf(c_0_44, negated_conjecture, (~p__d__subclass(c__NightTime,c__Entity)), inference(spm,[status(thm)],[c_0_35, c_0_36])).\\ncnf(c_0_45, plain, (p__d__subclass(c__NightTime,c__TimeInterval)), inference(split_conjunct,[status(thm)],[miloA2749])).\\ncnf(c_0_46, plain, ($false), inference(cdclpropres,[status(thm)],[c_0_37, c_0_38, c_0_39, c_0_40, c_0_41, c_0_42, c_0_43, c_0_44, c_0_45]), ['proof']).\\n# SZS output end CNFRefutation\\n# Training examples: 0 positive, 0 negative\\n\""
  ],
  "./goals/whitebox/whiteBoxTruthTest1315.adimen.tstp": [
    "ProverResult.PROOF_FOUND",
    "b\"# No SInE strategy applied\\n# Auto-Mode selected heuristic SAT001_MinMin_NN\\n# and selection function SelectMaxLComplexAvoidPosPred.\\n#\\n# Presaturation interreduction done\\n# SatCheck found unsatisfiable ground set\\n\\n# Proof found!\\n# SZS status Theorem\\n# SZS output start CNFRefutation\\nfof(whiteBoxTruthTest1315, conjecture, ?[X171]:~(p__d__instance(X171,c__Circle)), file('./goals/whitebox/whiteBoxTruthTest1315.adimen.tstp', whiteBoxTruthTest1315)).\\nfof(predefinitionsA12, axiom, ![X1, X2, X3]:((p__d__instance(X1,X2)&p__d__subclass(X2,X3))=>p__d__instance(X1,X3)), file('adimen.sumo.tstp', predefinitionsA12)).\\nfof(mergeA1778, axiom, p__d__subclass(c__Circle,c__Oval), file('adimen.sumo.tstp', mergeA1778)).\\nfof(predefinitionsA15, axiom, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~(p__d__instance(X6,X4))|~(p__d__instance(X6,X5)))), file('adimen.sumo.tstp', predefinitionsA15)).\\nfof(predefinitionsA24, axiom, ((((![X7, X8, X9]:(p__d__disjointDecomposition3(X7,X8,X9)<=>p__d__disjoint(X8,X9))&![X7, X8, X9, X10]:(p__d__disjointDecomposition4(X7,X8,X9,X10)<=>((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X9,X10))))&![X7, X8, X9, X10, X11]:(p__d__disjointDecomposition5(X7,X8,X9,X10,X11)<=>(((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X10,X11))))&![X7, X8, X9, X10, X11, X12]:(p__d__disjointDecomposition6(X7,X8,X9,X10,X11,X12)<=>(((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X11,X12))))&![X7, X8, X9, X10, X11, X12, X13]:(p__d__disjointDecomposition7(X7,X8,X9,X10,X11,X12,X13)<=>((((((((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X8,X13))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X9,X13))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X10,X13))&p__d__disjoint(X11,X12))&p__d__disjoint(X11,X13))&p__d__disjoint(X12,X13)))), file('adimen.sumo.tstp', predefinitionsA24)).\\nfof(mergeA1777, axiom, p__d__subclass(c__Oval,c__ClosedTwoDimensionalFigure), file('adimen.sumo.tstp', mergeA1777)).\\nfof(predefinitionsA8, axiom, ![X1, X2, X3]:((p__d__subclass(X1,X2)&p__d__subclass(X2,X3))=>p__d__subclass(X1,X3)), file('adimen.sumo.tstp', predefinitionsA8)).\\nfof(mergeA331, axiom, p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition), file('adimen.sumo.tstp', mergeA331)).\\nfof(mergeA575, axiom, p__d__subclass(c__Class,c__SetOrClass), file('adimen.sumo.tstp', mergeA575)).\\nfof(mergeA351, axiom, p__d__subclass(c__InternalAttribute,c__Attribute), file('adimen.sumo.tstp', mergeA351)).\\nfof(mergeA3615, axiom, p__d__subclass(c__ClosedTwoDimensionalFigure,c__TwoDimensionalFigure), file('adimen.sumo.tstp', mergeA3615)).\\nfof(mergeA3596, axiom, p__d__subclass(c__ShapeAttribute,c__StructureAttribute), file('adimen.sumo.tstp', mergeA3596)).\\nfof(mergeA3595, axiom, p__d__subclass(c__StructureAttribute,c__InternalAttribute), file('adimen.sumo.tstp', mergeA3595)).\\nfof(mergeA3607, axiom, p__d__subclass(c__GeometricFigure,c__ShapeAttribute), file('adimen.sumo.tstp', mergeA3607)).\\nfof(mergeA3611, axiom, p__d__subclass(c__TwoDimensionalFigure,c__GeometricFigure), file('adimen.sumo.tstp', mergeA3611)).\\nfof(mergeA333, axiom, p__d__subclass(c__Attribute,c__Abstract), file('adimen.sumo.tstp', mergeA333)).\\nfof(mergeA330, axiom, p__d__subclass(c__Abstract,c__Entity), file('adimen.sumo.tstp', mergeA330)).\\nfof(mergeA177, axiom, ![X7]:(p__d__instance(X7,c__Class)<=>p__d__subclass(X7,c__Entity)), file('adimen.sumo.tstp', mergeA177)).\\nfof(c_0_18, negated_conjecture, ~(?[X171]:~p__d__instance(X171,c__Circle)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[whiteBoxTruthTest1315])])).\\nfof(c_0_19, plain, ![X7776, X7777, X7778]:(~p__d__instance(X7776,X7777)|~p__d__subclass(X7777,X7778)|p__d__instance(X7776,X7778)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA12])])).\\nfof(c_0_20, negated_conjecture, ![X14607]:p__d__instance(X14607,c__Circle), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])).\\ncnf(c_0_21, plain, (p__d__instance(X1,X3)|~p__d__instance(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).\\ncnf(c_0_22, plain, (p__d__subclass(c__Circle,c__Oval)), inference(split_conjunct,[status(thm)],[mergeA1778])).\\ncnf(c_0_23, negated_conjecture, (p__d__instance(X1,c__Circle)), inference(split_conjunct,[status(thm)],[c_0_20])).\\nfof(c_0_24, plain, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~p__d__instance(X6,X4)|~p__d__instance(X6,X5))), inference(fof_simplification,[status(thm)],[predefinitionsA15])).\\nfof(c_0_25, plain, ![X7870, X7871, X7872, X7873, X7874, X7875, X7876, X7877, X7878, X7879, X7880, X7881, X7882, X7883, X7884, X7885, X7886, X7887, X7888, X7889, X7890, X7891, X7892, X7893, X7894]:((((((~p__d__disjointDecomposition3(X7870,X7871,X7872)|p__d__disjoint(X7871,X7872))&(~p__d__disjoint(X7871,X7872)|p__d__disjointDecomposition3(X7870,X7871,X7872)))&((((p__d__disjoint(X7874,X7875)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))&(p__d__disjoint(X7874,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(p__d__disjoint(X7875,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(~p__d__disjoint(X7874,X7875)|~p__d__disjoint(X7874,X7876)|~p__d__disjoint(X7875,X7876)|p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))))&(((((((p__d__disjoint(X7878,X7879)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))&(p__d__disjoint(X7878,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7878,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7880,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(~p__d__disjoint(X7878,X7879)|~p__d__disjoint(X7878,X7880)|~p__d__disjoint(X7878,X7881)|~p__d__disjoint(X7879,X7880)|~p__d__disjoint(X7879,X7881)|~p__d__disjoint(X7880,X7881)|p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))))&(((((((((((p__d__disjoint(X7883,X7884)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))&(p__d__disjoint(X7883,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7886,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(~p__d__disjoint(X7883,X7884)|~p__d__disjoint(X7883,X7885)|~p__d__disjoint(X7883,X7886)|~p__d__disjoint(X7883,X7887)|~p__d__disjoint(X7884,X7885)|~p__d__disjoint(X7884,X7886)|~p__d__disjoint(X7884,X7887)|~p__d__disjoint(X7885,X7886)|~p__d__disjoint(X7885,X7887)|~p__d__disjoint(X7886,X7887)|p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))))&((((((((((((((((p__d__disjoint(X7889,X7890)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894))&(p__d__disjoint(X7889,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7893,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(~p__d__disjoint(X7889,X7890)|~p__d__disjoint(X7889,X7891)|~p__d__disjoint(X7889,X7892)|~p__d__disjoint(X7889,X7893)|~p__d__disjoint(X7889,X7894)|~p__d__disjoint(X7890,X7891)|~p__d__disjoint(X7890,X7892)|~p__d__disjoint(X7890,X7893)|~p__d__disjoint(X7890,X7894)|~p__d__disjoint(X7891,X7892)|~p__d__disjoint(X7891,X7893)|~p__d__disjoint(X7891,X7894)|~p__d__disjoint(X7892,X7893)|~p__d__disjoint(X7892,X7894)|~p__d__disjoint(X7893,X7894)|p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA24])])])])).\\ncnf(c_0_26, plain, (p__d__subclass(c__Oval,c__ClosedTwoDimensionalFigure)), inference(split_conjunct,[status(thm)],[mergeA1777])).\\ncnf(c_0_27, plain, (p__d__instance(X1,c__Oval)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])])).\\nfof(c_0_28, plain, ![X7771, X7772, X7773]:(~p__d__subclass(X7771,X7772)|~p__d__subclass(X7772,X7773)|p__d__subclass(X7771,X7773)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA8])])).\\nfof(c_0_29, plain, ![X7779, X7780, X7781, X7782, X7783]:((~p__d__disjoint(X7779,X7780)|(~p__d__instance(X7781,X7779)|~p__d__instance(X7781,X7780)))&((p__d__instance(esk1_2(X7782,X7783),X7782)|p__d__disjoint(X7782,X7783))&(p__d__instance(esk1_2(X7782,X7783),X7783)|p__d__disjoint(X7782,X7783)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])])])])).\\ncnf(c_0_30, plain, (p__d__disjoint(X1,X2)|~p__d__disjointDecomposition6(X3,X4,X1,X2,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_25])).\\ncnf(c_0_31, plain, (p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition)), inference(split_conjunct,[status(thm)],[mergeA331])).\\ncnf(c_0_32, plain, (p__d__subclass(c__Class,c__SetOrClass)), inference(split_conjunct,[status(thm)],[mergeA575])).\\ncnf(c_0_33, plain, (p__d__subclass(c__InternalAttribute,c__Attribute)), inference(split_conjunct,[status(thm)],[mergeA351])).\\ncnf(c_0_34, plain, (p__d__subclass(c__ClosedTwoDimensionalFigure,c__TwoDimensionalFigure)), inference(split_conjunct,[status(thm)],[mergeA3615])).\\ncnf(c_0_35, plain, (p__d__instance(X1,c__ClosedTwoDimensionalFigure)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_26]), c_0_27])])).\\ncnf(c_0_36, plain, (p__d__subclass(c__ShapeAttribute,c__StructureAttribute)), inference(split_conjunct,[status(thm)],[mergeA3596])).\\ncnf(c_0_37, plain, (p__d__subclass(c__StructureAttribute,c__InternalAttribute)), inference(split_conjunct,[status(thm)],[mergeA3595])).\\ncnf(c_0_38, plain, (p__d__subclass(c__GeometricFigure,c__ShapeAttribute)), inference(split_conjunct,[status(thm)],[mergeA3607])).\\ncnf(c_0_39, plain, (p__d__subclass(c__TwoDimensionalFigure,c__GeometricFigure)), inference(split_conjunct,[status(thm)],[mergeA3611])).\\ncnf(c_0_40, plain, (p__d__subclass(X1,X3)|~p__d__subclass(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).\\ncnf(c_0_41, plain, (p__d__subclass(c__Attribute,c__Abstract)), inference(split_conjunct,[status(thm)],[mergeA333])).\\ncnf(c_0_42, plain, (p__d__subclass(c__Abstract,c__Entity)), inference(split_conjunct,[status(thm)],[mergeA330])).\\ncnf(c_0_43, plain, (~p__d__disjoint(X1,X2)|~p__d__instance(X3,X1)|~p__d__instance(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).\\ncnf(c_0_44, plain, (p__d__disjoint(c__Attribute,c__SetOrClass)), inference(spm,[status(thm)],[c_0_30, c_0_31])).\\nfof(c_0_45, plain, ![X8007]:((~p__d__instance(X8007,c__Class)|p__d__subclass(X8007,c__Entity))&(~p__d__subclass(X8007,c__Entity)|p__d__instance(X8007,c__Class))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mergeA177])])).\\ncnf(c_0_46, plain, (p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Class)), inference(spm,[status(thm)],[c_0_21, c_0_32])).\\ncnf(c_0_47, plain, (p__d__instance(X1,c__Attribute)|~p__d__instance(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_21, c_0_33])).\\ncnf(c_0_48, plain, (p__d__instance(X1,c__TwoDimensionalFigure)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_34]), c_0_35])])).\\ncnf(c_0_49, plain, (p__d__instance(X1,c__StructureAttribute)|~p__d__instance(X1,c__ShapeAttribute)), inference(spm,[status(thm)],[c_0_21, c_0_36])).\\ncnf(c_0_50, plain, (p__d__instance(X1,c__InternalAttribute)|~p__d__instance(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_21, c_0_37])).\\ncnf(c_0_51, plain, (p__d__instance(X1,c__ShapeAttribute)|~p__d__instance(X1,c__GeometricFigure)), inference(spm,[status(thm)],[c_0_21, c_0_38])).\\ncnf(c_0_52, plain, (p__d__instance(X1,c__GeometricFigure)|~p__d__instance(X1,c__TwoDimensionalFigure)), inference(spm,[status(thm)],[c_0_21, c_0_39])).\\ncnf(c_0_53, plain, (p__d__subclass(X1,c__Abstract)|~p__d__subclass(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_40, c_0_41])).\\ncnf(c_0_54, plain, (p__d__subclass(X1,c__Entity)|~p__d__subclass(X1,c__Abstract)), inference(spm,[status(thm)],[c_0_40, c_0_42])).\\ncnf(c_0_55, plain, (p__d__subclass(X1,c__Attribute)|~p__d__subclass(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_40, c_0_33])).\\ncnf(c_0_56, plain, (p__d__subclass(X1,c__ClosedTwoDimensionalFigure)|~p__d__subclass(X1,c__Oval)), inference(spm,[status(thm)],[c_0_40, c_0_26])).\\ncnf(c_0_57, plain, (p__d__subclass(X1,c__TwoDimensionalFigure)|~p__d__subclass(X1,c__ClosedTwoDimensionalFigure)), inference(spm,[status(thm)],[c_0_40, c_0_34])).\\ncnf(c_0_58, plain, (p__d__subclass(X1,c__StructureAttribute)|~p__d__subclass(X1,c__ShapeAttribute)), inference(spm,[status(thm)],[c_0_40, c_0_36])).\\ncnf(c_0_59, plain, (p__d__subclass(X1,c__InternalAttribute)|~p__d__subclass(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_40, c_0_37])).\\ncnf(c_0_60, plain, (p__d__subclass(X1,c__ShapeAttribute)|~p__d__subclass(X1,c__GeometricFigure)), inference(spm,[status(thm)],[c_0_40, c_0_38])).\\ncnf(c_0_61, plain, (p__d__subclass(X1,c__GeometricFigure)|~p__d__subclass(X1,c__TwoDimensionalFigure)), inference(spm,[status(thm)],[c_0_40, c_0_39])).\\ncnf(c_0_62, plain, (~p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_43, c_0_44])).\\ncnf(c_0_63, plain, (p__d__instance(X1,c__Class)|~p__d__subclass(X1,c__Entity)), inference(split_conjunct,[status(thm)],[c_0_45])).\\ncnf(c_0_64, plain, ($false), inference(cdclpropres,[status(thm)],[c_0_46, c_0_47, c_0_48, c_0_49, c_0_50, c_0_51, c_0_52, c_0_53, c_0_54, c_0_55, c_0_56, c_0_57, c_0_58, c_0_59, c_0_60, c_0_61, c_0_62, c_0_63, c_0_22]), ['proof']).\\n# SZS output end CNFRefutation\\n# Training examples: 0 positive, 0 negative\\n\""
  ],
  "./goals/whitebox/whiteBoxTruthTest1324.adimen.tstp": [
    "ProverResult.PROOF_FOUND",
    "b\"# No SInE strategy applied\\n# Auto-Mode selected heuristic SAT001_MinMin_NN\\n# and selection function SelectMaxLComplexAvoidPosPred.\\n#\\n# Presaturation interreduction done\\n# SatCheck found unsatisfiable ground set\\n\\n# Proof found!\\n# SZS status Theorem\\n# SZS output start CNFRefutation\\nfof(whiteBoxTruthTest1324, conjecture, ?[X174]:~(p__d__instance(X174,c__GeometricPoint)), file('./goals/whitebox/whiteBoxTruthTest1324.adimen.tstp', whiteBoxTruthTest1324)).\\nfof(predefinitionsA12, axiom, ![X1, X2, X3]:((p__d__instance(X1,X2)&p__d__subclass(X2,X3))=>p__d__instance(X1,X3)), file('adimen.sumo.tstp', predefinitionsA12)).\\nfof(mergeA3609, axiom, p__d__subclass(c__GeometricPoint,c__GeometricFigure), file('adimen.sumo.tstp', mergeA3609)).\\nfof(predefinitionsA15, axiom, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~(p__d__instance(X6,X4))|~(p__d__instance(X6,X5)))), file('adimen.sumo.tstp', predefinitionsA15)).\\nfof(predefinitionsA24, axiom, ((((![X7, X8, X9]:(p__d__disjointDecomposition3(X7,X8,X9)<=>p__d__disjoint(X8,X9))&![X7, X8, X9, X10]:(p__d__disjointDecomposition4(X7,X8,X9,X10)<=>((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X9,X10))))&![X7, X8, X9, X10, X11]:(p__d__disjointDecomposition5(X7,X8,X9,X10,X11)<=>(((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X10,X11))))&![X7, X8, X9, X10, X11, X12]:(p__d__disjointDecomposition6(X7,X8,X9,X10,X11,X12)<=>(((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X11,X12))))&![X7, X8, X9, X10, X11, X12, X13]:(p__d__disjointDecomposition7(X7,X8,X9,X10,X11,X12,X13)<=>((((((((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X8,X13))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X9,X13))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X10,X13))&p__d__disjoint(X11,X12))&p__d__disjoint(X11,X13))&p__d__disjoint(X12,X13)))), file('adimen.sumo.tstp', predefinitionsA24)).\\nfof(predefinitionsA8, axiom, ![X1, X2, X3]:((p__d__subclass(X1,X2)&p__d__subclass(X2,X3))=>p__d__subclass(X1,X3)), file('adimen.sumo.tstp', predefinitionsA8)).\\nfof(mergeA3607, axiom, p__d__subclass(c__GeometricFigure,c__ShapeAttribute), file('adimen.sumo.tstp', mergeA3607)).\\nfof(mergeA331, axiom, p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition), file('adimen.sumo.tstp', mergeA331)).\\nfof(mergeA333, axiom, p__d__subclass(c__Attribute,c__Abstract), file('adimen.sumo.tstp', mergeA333)).\\nfof(mergeA330, axiom, p__d__subclass(c__Abstract,c__Entity), file('adimen.sumo.tstp', mergeA330)).\\nfof(mergeA351, axiom, p__d__subclass(c__InternalAttribute,c__Attribute), file('adimen.sumo.tstp', mergeA351)).\\nfof(mergeA3596, axiom, p__d__subclass(c__ShapeAttribute,c__StructureAttribute), file('adimen.sumo.tstp', mergeA3596)).\\nfof(mergeA3595, axiom, p__d__subclass(c__StructureAttribute,c__InternalAttribute), file('adimen.sumo.tstp', mergeA3595)).\\nfof(mergeA575, axiom, p__d__subclass(c__Class,c__SetOrClass), file('adimen.sumo.tstp', mergeA575)).\\nfof(mergeA177, axiom, ![X7]:(p__d__instance(X7,c__Class)<=>p__d__subclass(X7,c__Entity)), file('adimen.sumo.tstp', mergeA177)).\\nfof(c_0_15, negated_conjecture, ~(?[X174]:~p__d__instance(X174,c__GeometricPoint)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[whiteBoxTruthTest1324])])).\\nfof(c_0_16, plain, ![X7776, X7777, X7778]:(~p__d__instance(X7776,X7777)|~p__d__subclass(X7777,X7778)|p__d__instance(X7776,X7778)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA12])])).\\nfof(c_0_17, negated_conjecture, ![X14607]:p__d__instance(X14607,c__GeometricPoint), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).\\ncnf(c_0_18, plain, (p__d__instance(X1,X3)|~p__d__instance(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).\\ncnf(c_0_19, plain, (p__d__subclass(c__GeometricPoint,c__GeometricFigure)), inference(split_conjunct,[status(thm)],[mergeA3609])).\\ncnf(c_0_20, negated_conjecture, (p__d__instance(X1,c__GeometricPoint)), inference(split_conjunct,[status(thm)],[c_0_17])).\\nfof(c_0_21, plain, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~p__d__instance(X6,X4)|~p__d__instance(X6,X5))), inference(fof_simplification,[status(thm)],[predefinitionsA15])).\\nfof(c_0_22, plain, ![X7870, X7871, X7872, X7873, X7874, X7875, X7876, X7877, X7878, X7879, X7880, X7881, X7882, X7883, X7884, X7885, X7886, X7887, X7888, X7889, X7890, X7891, X7892, X7893, X7894]:((((((~p__d__disjointDecomposition3(X7870,X7871,X7872)|p__d__disjoint(X7871,X7872))&(~p__d__disjoint(X7871,X7872)|p__d__disjointDecomposition3(X7870,X7871,X7872)))&((((p__d__disjoint(X7874,X7875)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))&(p__d__disjoint(X7874,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(p__d__disjoint(X7875,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(~p__d__disjoint(X7874,X7875)|~p__d__disjoint(X7874,X7876)|~p__d__disjoint(X7875,X7876)|p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))))&(((((((p__d__disjoint(X7878,X7879)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))&(p__d__disjoint(X7878,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7878,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7880,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(~p__d__disjoint(X7878,X7879)|~p__d__disjoint(X7878,X7880)|~p__d__disjoint(X7878,X7881)|~p__d__disjoint(X7879,X7880)|~p__d__disjoint(X7879,X7881)|~p__d__disjoint(X7880,X7881)|p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))))&(((((((((((p__d__disjoint(X7883,X7884)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))&(p__d__disjoint(X7883,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7886,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(~p__d__disjoint(X7883,X7884)|~p__d__disjoint(X7883,X7885)|~p__d__disjoint(X7883,X7886)|~p__d__disjoint(X7883,X7887)|~p__d__disjoint(X7884,X7885)|~p__d__disjoint(X7884,X7886)|~p__d__disjoint(X7884,X7887)|~p__d__disjoint(X7885,X7886)|~p__d__disjoint(X7885,X7887)|~p__d__disjoint(X7886,X7887)|p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))))&((((((((((((((((p__d__disjoint(X7889,X7890)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894))&(p__d__disjoint(X7889,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7893,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(~p__d__disjoint(X7889,X7890)|~p__d__disjoint(X7889,X7891)|~p__d__disjoint(X7889,X7892)|~p__d__disjoint(X7889,X7893)|~p__d__disjoint(X7889,X7894)|~p__d__disjoint(X7890,X7891)|~p__d__disjoint(X7890,X7892)|~p__d__disjoint(X7890,X7893)|~p__d__disjoint(X7890,X7894)|~p__d__disjoint(X7891,X7892)|~p__d__disjoint(X7891,X7893)|~p__d__disjoint(X7891,X7894)|~p__d__disjoint(X7892,X7893)|~p__d__disjoint(X7892,X7894)|~p__d__disjoint(X7893,X7894)|p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA24])])])])).\\nfof(c_0_23, plain, ![X7771, X7772, X7773]:(~p__d__subclass(X7771,X7772)|~p__d__subclass(X7772,X7773)|p__d__subclass(X7771,X7773)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA8])])).\\ncnf(c_0_24, plain, (p__d__subclass(c__GeometricFigure,c__ShapeAttribute)), inference(split_conjunct,[status(thm)],[mergeA3607])).\\ncnf(c_0_25, plain, (p__d__instance(X1,c__GeometricFigure)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])])).\\nfof(c_0_26, plain, ![X7779, X7780, X7781, X7782, X7783]:((~p__d__disjoint(X7779,X7780)|(~p__d__instance(X7781,X7779)|~p__d__instance(X7781,X7780)))&((p__d__instance(esk1_2(X7782,X7783),X7782)|p__d__disjoint(X7782,X7783))&(p__d__instance(esk1_2(X7782,X7783),X7783)|p__d__disjoint(X7782,X7783)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])])).\\ncnf(c_0_27, plain, (p__d__disjoint(X1,X2)|~p__d__disjointDecomposition6(X3,X4,X1,X2,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_22])).\\ncnf(c_0_28, plain, (p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition)), inference(split_conjunct,[status(thm)],[mergeA331])).\\ncnf(c_0_29, plain, (p__d__subclass(X1,X3)|~p__d__subclass(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).\\ncnf(c_0_30, plain, (p__d__subclass(c__Attribute,c__Abstract)), inference(split_conjunct,[status(thm)],[mergeA333])).\\ncnf(c_0_31, plain, (p__d__subclass(c__Abstract,c__Entity)), inference(split_conjunct,[status(thm)],[mergeA330])).\\ncnf(c_0_32, plain, (p__d__subclass(c__InternalAttribute,c__Attribute)), inference(split_conjunct,[status(thm)],[mergeA351])).\\ncnf(c_0_33, plain, (p__d__subclass(c__ShapeAttribute,c__StructureAttribute)), inference(split_conjunct,[status(thm)],[mergeA3596])).\\ncnf(c_0_34, plain, (p__d__subclass(c__StructureAttribute,c__InternalAttribute)), inference(split_conjunct,[status(thm)],[mergeA3595])).\\ncnf(c_0_35, plain, (p__d__subclass(c__Class,c__SetOrClass)), inference(split_conjunct,[status(thm)],[mergeA575])).\\ncnf(c_0_36, plain, (p__d__instance(X1,c__ShapeAttribute)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_24]), c_0_25])])).\\ncnf(c_0_37, plain, (~p__d__disjoint(X1,X2)|~p__d__instance(X3,X1)|~p__d__instance(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).\\ncnf(c_0_38, plain, (p__d__disjoint(c__Attribute,c__SetOrClass)), inference(spm,[status(thm)],[c_0_27, c_0_28])).\\nfof(c_0_39, plain, ![X8007]:((~p__d__instance(X8007,c__Class)|p__d__subclass(X8007,c__Entity))&(~p__d__subclass(X8007,c__Entity)|p__d__instance(X8007,c__Class))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mergeA177])])).\\ncnf(c_0_40, plain, (p__d__subclass(X1,c__Abstract)|~p__d__subclass(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_29, c_0_30])).\\ncnf(c_0_41, plain, (p__d__subclass(X1,c__Entity)|~p__d__subclass(X1,c__Abstract)), inference(spm,[status(thm)],[c_0_29, c_0_31])).\\ncnf(c_0_42, plain, (p__d__subclass(X1,c__Attribute)|~p__d__subclass(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_29, c_0_32])).\\ncnf(c_0_43, plain, (p__d__subclass(X1,c__StructureAttribute)|~p__d__subclass(X1,c__ShapeAttribute)), inference(spm,[status(thm)],[c_0_29, c_0_33])).\\ncnf(c_0_44, plain, (p__d__subclass(X1,c__InternalAttribute)|~p__d__subclass(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_29, c_0_34])).\\ncnf(c_0_45, plain, (p__d__subclass(X1,c__ShapeAttribute)|~p__d__subclass(X1,c__GeometricFigure)), inference(spm,[status(thm)],[c_0_29, c_0_24])).\\ncnf(c_0_46, plain, (p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Class)), inference(spm,[status(thm)],[c_0_18, c_0_35])).\\ncnf(c_0_47, plain, (p__d__instance(X1,c__Attribute)|~p__d__instance(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_18, c_0_32])).\\ncnf(c_0_48, plain, (p__d__instance(X1,c__StructureAttribute)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_33]), c_0_36])])).\\ncnf(c_0_49, plain, (p__d__instance(X1,c__InternalAttribute)|~p__d__instance(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_18, c_0_34])).\\ncnf(c_0_50, plain, (~p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_37, c_0_38])).\\ncnf(c_0_51, plain, (p__d__instance(X1,c__Class)|~p__d__subclass(X1,c__Entity)), inference(split_conjunct,[status(thm)],[c_0_39])).\\ncnf(c_0_52, plain, ($false), inference(cdclpropres,[status(thm)],[c_0_40, c_0_41, c_0_42, c_0_43, c_0_44, c_0_45, c_0_46, c_0_47, c_0_48, c_0_49, c_0_50, c_0_51, c_0_19]), ['proof']).\\n# SZS output end CNFRefutation\\n# Training examples: 0 positive, 0 negative\\n\""
  ],
  "./goals/whitebox/whiteBoxTruthTest4290.adimen.tstp": [
    "ProverResult.PROOF_FOUND",
    "b\"# No SInE strategy applied\\n# Auto-Mode selected heuristic SAT001_MinMin_NN\\n# and selection function SelectMaxLComplexAvoidPosPred.\\n#\\n# Presaturation interreduction done\\n# SatCheck found unsatisfiable ground set\\n\\n# Proof found!\\n# SZS status Theorem\\n# SZS output start CNFRefutation\\nfof(whiteBoxTruthTest4290, conjecture, ?[X272]:p__d__instance(X272,c__RightTriangle), file('./goals/whitebox/whiteBoxTruthTest4290.adimen.tstp', whiteBoxTruthTest4290)).\\nfof(predefinitionsA8, axiom, ![X1, X2, X3]:((p__d__subclass(X1,X2)&p__d__subclass(X2,X3))=>p__d__subclass(X1,X3)), file('adimen.sumo.tstp', predefinitionsA8)).\\nfof(mergeA176, axiom, ![X7]:(p__d__subclass(X7,c__Entity)=>?[X35]:p__d__instance(X35,X7)), file('adimen.sumo.tstp', mergeA176)).\\nfof(mergeA333, axiom, p__d__subclass(c__Attribute,c__Abstract), file('adimen.sumo.tstp', mergeA333)).\\nfof(mergeA330, axiom, p__d__subclass(c__Abstract,c__Entity), file('adimen.sumo.tstp', mergeA330)).\\nfof(mergeA351, axiom, p__d__subclass(c__InternalAttribute,c__Attribute), file('adimen.sumo.tstp', mergeA351)).\\nfof(mergeA3615, axiom, p__d__subclass(c__ClosedTwoDimensionalFigure,c__TwoDimensionalFigure), file('adimen.sumo.tstp', mergeA3615)).\\nfof(mergeA3596, axiom, p__d__subclass(c__ShapeAttribute,c__StructureAttribute), file('adimen.sumo.tstp', mergeA3596)).\\nfof(mergeA3595, axiom, p__d__subclass(c__StructureAttribute,c__InternalAttribute), file('adimen.sumo.tstp', mergeA3595)).\\nfof(mergeA3607, axiom, p__d__subclass(c__GeometricFigure,c__ShapeAttribute), file('adimen.sumo.tstp', mergeA3607)).\\nfof(mergeA3611, axiom, p__d__subclass(c__TwoDimensionalFigure,c__GeometricFigure), file('adimen.sumo.tstp', mergeA3611)).\\nfof(miloA759, axiom, p__d__subclass(c__Polygon,c__ClosedTwoDimensionalFigure), file('adimen.sumo.tstp', miloA759)).\\nfof(miloA761, axiom, p__d__subclass(c__Triangle,c__Polygon), file('adimen.sumo.tstp', miloA761)).\\nfof(miloA763, axiom, p__d__subclass(c__RightTriangle,c__Triangle), file('adimen.sumo.tstp', miloA763)).\\nfof(c_0_14, negated_conjecture, ~(?[X272]:p__d__instance(X272,c__RightTriangle)), inference(assume_negation,[status(cth)],[whiteBoxTruthTest4290])).\\nfof(c_0_15, plain, ![X7771, X7772, X7773]:(~p__d__subclass(X7771,X7772)|~p__d__subclass(X7772,X7773)|p__d__subclass(X7771,X7773)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA8])])).\\nfof(c_0_16, negated_conjecture, ![X14607]:~p__d__instance(X14607,c__RightTriangle), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).\\nfof(c_0_17, plain, ![X8005]:(~p__d__subclass(X8005,c__Entity)|p__d__instance(esk13_1(X8005),X8005)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mergeA176])])])).\\ncnf(c_0_18, plain, (p__d__subclass(X1,X3)|~p__d__subclass(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).\\ncnf(c_0_19, plain, (p__d__subclass(c__Attribute,c__Abstract)), inference(split_conjunct,[status(thm)],[mergeA333])).\\ncnf(c_0_20, plain, (p__d__subclass(c__Abstract,c__Entity)), inference(split_conjunct,[status(thm)],[mergeA330])).\\ncnf(c_0_21, plain, (p__d__subclass(c__InternalAttribute,c__Attribute)), inference(split_conjunct,[status(thm)],[mergeA351])).\\ncnf(c_0_22, plain, (p__d__subclass(c__ClosedTwoDimensionalFigure,c__TwoDimensionalFigure)), inference(split_conjunct,[status(thm)],[mergeA3615])).\\ncnf(c_0_23, plain, (p__d__subclass(c__ShapeAttribute,c__StructureAttribute)), inference(split_conjunct,[status(thm)],[mergeA3596])).\\ncnf(c_0_24, plain, (p__d__subclass(c__StructureAttribute,c__InternalAttribute)), inference(split_conjunct,[status(thm)],[mergeA3595])).\\ncnf(c_0_25, plain, (p__d__subclass(c__GeometricFigure,c__ShapeAttribute)), inference(split_conjunct,[status(thm)],[mergeA3607])).\\ncnf(c_0_26, plain, (p__d__subclass(c__TwoDimensionalFigure,c__GeometricFigure)), inference(split_conjunct,[status(thm)],[mergeA3611])).\\ncnf(c_0_27, plain, (p__d__subclass(c__Polygon,c__ClosedTwoDimensionalFigure)), inference(split_conjunct,[status(thm)],[miloA759])).\\ncnf(c_0_28, plain, (p__d__subclass(c__Triangle,c__Polygon)), inference(split_conjunct,[status(thm)],[miloA761])).\\ncnf(c_0_29, negated_conjecture, (~p__d__instance(X1,c__RightTriangle)), inference(split_conjunct,[status(thm)],[c_0_16])).\\ncnf(c_0_30, plain, (p__d__instance(esk13_1(X1),X1)|~p__d__subclass(X1,c__Entity)), inference(split_conjunct,[status(thm)],[c_0_17])).\\ncnf(c_0_31, plain, (p__d__subclass(X1,c__Abstract)|~p__d__subclass(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_18, c_0_19])).\\ncnf(c_0_32, plain, (p__d__subclass(X1,c__Entity)|~p__d__subclass(X1,c__Abstract)), inference(spm,[status(thm)],[c_0_18, c_0_20])).\\ncnf(c_0_33, plain, (p__d__subclass(X1,c__Attribute)|~p__d__subclass(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_18, c_0_21])).\\ncnf(c_0_34, plain, (p__d__subclass(X1,c__TwoDimensionalFigure)|~p__d__subclass(X1,c__ClosedTwoDimensionalFigure)), inference(spm,[status(thm)],[c_0_18, c_0_22])).\\ncnf(c_0_35, plain, (p__d__subclass(X1,c__StructureAttribute)|~p__d__subclass(X1,c__ShapeAttribute)), inference(spm,[status(thm)],[c_0_18, c_0_23])).\\ncnf(c_0_36, plain, (p__d__subclass(X1,c__InternalAttribute)|~p__d__subclass(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_18, c_0_24])).\\ncnf(c_0_37, plain, (p__d__subclass(X1,c__ShapeAttribute)|~p__d__subclass(X1,c__GeometricFigure)), inference(spm,[status(thm)],[c_0_18, c_0_25])).\\ncnf(c_0_38, plain, (p__d__subclass(X1,c__GeometricFigure)|~p__d__subclass(X1,c__TwoDimensionalFigure)), inference(spm,[status(thm)],[c_0_18, c_0_26])).\\ncnf(c_0_39, plain, (p__d__subclass(X1,c__ClosedTwoDimensionalFigure)|~p__d__subclass(X1,c__Polygon)), inference(spm,[status(thm)],[c_0_18, c_0_27])).\\ncnf(c_0_40, plain, (p__d__subclass(X1,c__Polygon)|~p__d__subclass(X1,c__Triangle)), inference(spm,[status(thm)],[c_0_18, c_0_28])).\\ncnf(c_0_41, negated_conjecture, (~p__d__subclass(c__RightTriangle,c__Entity)), inference(spm,[status(thm)],[c_0_29, c_0_30])).\\ncnf(c_0_42, plain, (p__d__subclass(c__RightTriangle,c__Triangle)), inference(split_conjunct,[status(thm)],[miloA763])).\\ncnf(c_0_43, plain, ($false), inference(cdclpropres,[status(thm)],[c_0_31, c_0_32, c_0_33, c_0_34, c_0_35, c_0_36, c_0_37, c_0_38, c_0_39, c_0_40, c_0_41, c_0_42]), ['proof']).\\n# SZS output end CNFRefutation\\n# Training examples: 0 positive, 0 negative\\n\""
  ],
  "./goals/whitebox/whiteBoxTruthTest1250.adimen.tstp": [
    "ProverResult.PROOF_FOUND",
    "b\"# No SInE strategy applied\\n# Auto-Mode selected heuristic SAT001_MinMin_NN\\n# and selection function SelectMaxLComplexAvoidPosPred.\\n#\\n# Presaturation interreduction done\\n# SatCheck found unsatisfiable ground set\\n\\n# Proof found!\\n# SZS status Theorem\\n# SZS output start CNFRefutation\\nfof(whiteBoxTruthTest1250, conjecture, ?[X159]:~(p__d__instance(X159,c__FrequencyMeasure)), file('./goals/whitebox/whiteBoxTruthTest1250.adimen.tstp', whiteBoxTruthTest1250)).\\nfof(predefinitionsA12, axiom, ![X1, X2, X3]:((p__d__instance(X1,X2)&p__d__subclass(X2,X3))=>p__d__instance(X1,X3)), file('adimen.sumo.tstp', predefinitionsA12)).\\nfof(predefinitionsA15, axiom, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~(p__d__instance(X6,X4))|~(p__d__instance(X6,X5)))), file('adimen.sumo.tstp', predefinitionsA15)).\\nfof(predefinitionsA24, axiom, ((((![X7, X8, X9]:(p__d__disjointDecomposition3(X7,X8,X9)<=>p__d__disjoint(X8,X9))&![X7, X8, X9, X10]:(p__d__disjointDecomposition4(X7,X8,X9,X10)<=>((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X9,X10))))&![X7, X8, X9, X10, X11]:(p__d__disjointDecomposition5(X7,X8,X9,X10,X11)<=>(((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X10,X11))))&![X7, X8, X9, X10, X11, X12]:(p__d__disjointDecomposition6(X7,X8,X9,X10,X11,X12)<=>(((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X11,X12))))&![X7, X8, X9, X10, X11, X12, X13]:(p__d__disjointDecomposition7(X7,X8,X9,X10,X11,X12,X13)<=>((((((((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X8,X13))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X9,X13))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X10,X13))&p__d__disjoint(X11,X12))&p__d__disjoint(X11,X13))&p__d__disjoint(X12,X13)))), file('adimen.sumo.tstp', predefinitionsA24)).\\nfof(mergeA1652, axiom, p__d__subclass(c__FrequencyMeasure,c__TimeDependentQuantity), file('adimen.sumo.tstp', mergeA1652)).\\nfof(mergeA331, axiom, p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition), file('adimen.sumo.tstp', mergeA331)).\\nfof(predefinitionsA8, axiom, ![X1, X2, X3]:((p__d__subclass(X1,X2)&p__d__subclass(X2,X3))=>p__d__subclass(X1,X3)), file('adimen.sumo.tstp', predefinitionsA8)).\\nfof(mergeA572, axiom, p__d__subclass(c__TimeDependentQuantity,c__UnaryConstantFunctionQuantity), file('adimen.sumo.tstp', mergeA572)).\\nfof(mergeA330, axiom, p__d__subclass(c__Abstract,c__Entity), file('adimen.sumo.tstp', mergeA330)).\\nfof(mergeA332, axiom, p__d__subclass(c__Quantity,c__Abstract), file('adimen.sumo.tstp', mergeA332)).\\nfof(mergeA561, axiom, p__d__subclass(c__PhysicalQuantity,c__Quantity), file('adimen.sumo.tstp', mergeA561)).\\nfof(mergeA570, axiom, p__d__subclass(c__FunctionQuantity,c__PhysicalQuantity), file('adimen.sumo.tstp', mergeA570)).\\nfof(mergeA571, axiom, p__d__subclass(c__UnaryConstantFunctionQuantity,c__FunctionQuantity), file('adimen.sumo.tstp', mergeA571)).\\nfof(mergeA575, axiom, p__d__subclass(c__Class,c__SetOrClass), file('adimen.sumo.tstp', mergeA575)).\\nfof(mergeA177, axiom, ![X7]:(p__d__instance(X7,c__Class)<=>p__d__subclass(X7,c__Entity)), file('adimen.sumo.tstp', mergeA177)).\\nfof(c_0_15, negated_conjecture, ~(?[X159]:~p__d__instance(X159,c__FrequencyMeasure)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[whiteBoxTruthTest1250])])).\\nfof(c_0_16, plain, ![X7776, X7777, X7778]:(~p__d__instance(X7776,X7777)|~p__d__subclass(X7777,X7778)|p__d__instance(X7776,X7778)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA12])])).\\nfof(c_0_17, negated_conjecture, ![X14607]:p__d__instance(X14607,c__FrequencyMeasure), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).\\nfof(c_0_18, plain, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~p__d__instance(X6,X4)|~p__d__instance(X6,X5))), inference(fof_simplification,[status(thm)],[predefinitionsA15])).\\nfof(c_0_19, plain, ![X7870, X7871, X7872, X7873, X7874, X7875, X7876, X7877, X7878, X7879, X7880, X7881, X7882, X7883, X7884, X7885, X7886, X7887, X7888, X7889, X7890, X7891, X7892, X7893, X7894]:((((((~p__d__disjointDecomposition3(X7870,X7871,X7872)|p__d__disjoint(X7871,X7872))&(~p__d__disjoint(X7871,X7872)|p__d__disjointDecomposition3(X7870,X7871,X7872)))&((((p__d__disjoint(X7874,X7875)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))&(p__d__disjoint(X7874,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(p__d__disjoint(X7875,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(~p__d__disjoint(X7874,X7875)|~p__d__disjoint(X7874,X7876)|~p__d__disjoint(X7875,X7876)|p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))))&(((((((p__d__disjoint(X7878,X7879)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))&(p__d__disjoint(X7878,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7878,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7880,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(~p__d__disjoint(X7878,X7879)|~p__d__disjoint(X7878,X7880)|~p__d__disjoint(X7878,X7881)|~p__d__disjoint(X7879,X7880)|~p__d__disjoint(X7879,X7881)|~p__d__disjoint(X7880,X7881)|p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))))&(((((((((((p__d__disjoint(X7883,X7884)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))&(p__d__disjoint(X7883,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7886,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(~p__d__disjoint(X7883,X7884)|~p__d__disjoint(X7883,X7885)|~p__d__disjoint(X7883,X7886)|~p__d__disjoint(X7883,X7887)|~p__d__disjoint(X7884,X7885)|~p__d__disjoint(X7884,X7886)|~p__d__disjoint(X7884,X7887)|~p__d__disjoint(X7885,X7886)|~p__d__disjoint(X7885,X7887)|~p__d__disjoint(X7886,X7887)|p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))))&((((((((((((((((p__d__disjoint(X7889,X7890)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894))&(p__d__disjoint(X7889,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7893,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(~p__d__disjoint(X7889,X7890)|~p__d__disjoint(X7889,X7891)|~p__d__disjoint(X7889,X7892)|~p__d__disjoint(X7889,X7893)|~p__d__disjoint(X7889,X7894)|~p__d__disjoint(X7890,X7891)|~p__d__disjoint(X7890,X7892)|~p__d__disjoint(X7890,X7893)|~p__d__disjoint(X7890,X7894)|~p__d__disjoint(X7891,X7892)|~p__d__disjoint(X7891,X7893)|~p__d__disjoint(X7891,X7894)|~p__d__disjoint(X7892,X7893)|~p__d__disjoint(X7892,X7894)|~p__d__disjoint(X7893,X7894)|p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA24])])])])).\\ncnf(c_0_20, plain, (p__d__instance(X1,X3)|~p__d__instance(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).\\ncnf(c_0_21, plain, (p__d__subclass(c__FrequencyMeasure,c__TimeDependentQuantity)), inference(split_conjunct,[status(thm)],[mergeA1652])).\\ncnf(c_0_22, negated_conjecture, (p__d__instance(X1,c__FrequencyMeasure)), inference(split_conjunct,[status(thm)],[c_0_17])).\\nfof(c_0_23, plain, ![X7779, X7780, X7781, X7782, X7783]:((~p__d__disjoint(X7779,X7780)|(~p__d__instance(X7781,X7779)|~p__d__instance(X7781,X7780)))&((p__d__instance(esk1_2(X7782,X7783),X7782)|p__d__disjoint(X7782,X7783))&(p__d__instance(esk1_2(X7782,X7783),X7783)|p__d__disjoint(X7782,X7783)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])])])).\\ncnf(c_0_24, plain, (p__d__disjoint(X1,X2)|~p__d__disjointDecomposition6(X3,X1,X4,X2,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_19])).\\ncnf(c_0_25, plain, (p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition)), inference(split_conjunct,[status(thm)],[mergeA331])).\\nfof(c_0_26, plain, ![X7771, X7772, X7773]:(~p__d__subclass(X7771,X7772)|~p__d__subclass(X7772,X7773)|p__d__subclass(X7771,X7773)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA8])])).\\ncnf(c_0_27, plain, (p__d__subclass(c__TimeDependentQuantity,c__UnaryConstantFunctionQuantity)), inference(split_conjunct,[status(thm)],[mergeA572])).\\ncnf(c_0_28, plain, (p__d__instance(X1,c__TimeDependentQuantity)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])])).\\ncnf(c_0_29, plain, (~p__d__disjoint(X1,X2)|~p__d__instance(X3,X1)|~p__d__instance(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).\\ncnf(c_0_30, plain, (p__d__disjoint(c__Quantity,c__SetOrClass)), inference(spm,[status(thm)],[c_0_24, c_0_25])).\\ncnf(c_0_31, plain, (p__d__subclass(X1,X3)|~p__d__subclass(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).\\ncnf(c_0_32, plain, (p__d__subclass(c__Abstract,c__Entity)), inference(split_conjunct,[status(thm)],[mergeA330])).\\ncnf(c_0_33, plain, (p__d__subclass(c__Quantity,c__Abstract)), inference(split_conjunct,[status(thm)],[mergeA332])).\\ncnf(c_0_34, plain, (p__d__subclass(c__PhysicalQuantity,c__Quantity)), inference(split_conjunct,[status(thm)],[mergeA561])).\\ncnf(c_0_35, plain, (p__d__subclass(c__FunctionQuantity,c__PhysicalQuantity)), inference(split_conjunct,[status(thm)],[mergeA570])).\\ncnf(c_0_36, plain, (p__d__subclass(c__UnaryConstantFunctionQuantity,c__FunctionQuantity)), inference(split_conjunct,[status(thm)],[mergeA571])).\\ncnf(c_0_37, plain, (p__d__subclass(c__Class,c__SetOrClass)), inference(split_conjunct,[status(thm)],[mergeA575])).\\ncnf(c_0_38, plain, (p__d__instance(X1,c__UnaryConstantFunctionQuantity)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_27]), c_0_28])])).\\nfof(c_0_39, plain, ![X8007]:((~p__d__instance(X8007,c__Class)|p__d__subclass(X8007,c__Entity))&(~p__d__subclass(X8007,c__Entity)|p__d__instance(X8007,c__Class))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mergeA177])])).\\ncnf(c_0_40, plain, (~p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Quantity)), inference(spm,[status(thm)],[c_0_29, c_0_30])).\\ncnf(c_0_41, plain, (p__d__subclass(X1,c__Entity)|~p__d__subclass(X1,c__Abstract)), inference(spm,[status(thm)],[c_0_31, c_0_32])).\\ncnf(c_0_42, plain, (p__d__subclass(X1,c__Abstract)|~p__d__subclass(X1,c__Quantity)), inference(spm,[status(thm)],[c_0_31, c_0_33])).\\ncnf(c_0_43, plain, (p__d__subclass(X1,c__Quantity)|~p__d__subclass(X1,c__PhysicalQuantity)), inference(spm,[status(thm)],[c_0_31, c_0_34])).\\ncnf(c_0_44, plain, (p__d__subclass(X1,c__PhysicalQuantity)|~p__d__subclass(X1,c__FunctionQuantity)), inference(spm,[status(thm)],[c_0_31, c_0_35])).\\ncnf(c_0_45, plain, (p__d__subclass(X1,c__FunctionQuantity)|~p__d__subclass(X1,c__UnaryConstantFunctionQuantity)), inference(spm,[status(thm)],[c_0_31, c_0_36])).\\ncnf(c_0_46, plain, (p__d__subclass(X1,c__UnaryConstantFunctionQuantity)|~p__d__subclass(X1,c__TimeDependentQuantity)), inference(spm,[status(thm)],[c_0_31, c_0_27])).\\ncnf(c_0_47, plain, (p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Class)), inference(spm,[status(thm)],[c_0_20, c_0_37])).\\ncnf(c_0_48, plain, (p__d__instance(X1,c__Quantity)|~p__d__instance(X1,c__PhysicalQuantity)), inference(spm,[status(thm)],[c_0_20, c_0_34])).\\ncnf(c_0_49, plain, (p__d__instance(X1,c__PhysicalQuantity)|~p__d__instance(X1,c__FunctionQuantity)), inference(spm,[status(thm)],[c_0_20, c_0_35])).\\ncnf(c_0_50, plain, (p__d__instance(X1,c__FunctionQuantity)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_36]), c_0_38])])).\\ncnf(c_0_51, plain, (p__d__instance(X1,c__Class)|~p__d__subclass(X1,c__Entity)), inference(split_conjunct,[status(thm)],[c_0_39])).\\ncnf(c_0_52, plain, ($false), inference(cdclpropres,[status(thm)],[c_0_40, c_0_41, c_0_42, c_0_43, c_0_44, c_0_45, c_0_46, c_0_47, c_0_48, c_0_49, c_0_50, c_0_51, c_0_21]), ['proof']).\\n# SZS output end CNFRefutation\\n# Training examples: 0 positive, 0 negative\\n\""
  ],
  "./goals/whitebox/whiteBoxTruthTest4282.adimen.tstp": [
    "ProverResult.PROOF_FOUND",
    "b\"# No SInE strategy applied\\n# Auto-Mode selected heuristic SAT001_MinMin_NN\\n# and selection function SelectMaxLComplexAvoidPosPred.\\n#\\n# Presaturation interreduction done\\n# SatCheck found unsatisfiable ground set\\n\\n# Proof found!\\n# SZS status Theorem\\n# SZS output start CNFRefutation\\nfof(whiteBoxTruthTest4282, conjecture, ?[X610]:~(p__d__instance(X610,c__Polygon)), file('./goals/whitebox/whiteBoxTruthTest4282.adimen.tstp', whiteBoxTruthTest4282)).\\nfof(predefinitionsA12, axiom, ![X1, X2, X3]:((p__d__instance(X1,X2)&p__d__subclass(X2,X3))=>p__d__instance(X1,X3)), file('adimen.sumo.tstp', predefinitionsA12)).\\nfof(miloA759, axiom, p__d__subclass(c__Polygon,c__ClosedTwoDimensionalFigure), file('adimen.sumo.tstp', miloA759)).\\nfof(predefinitionsA15, axiom, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~(p__d__instance(X6,X4))|~(p__d__instance(X6,X5)))), file('adimen.sumo.tstp', predefinitionsA15)).\\nfof(predefinitionsA24, axiom, ((((![X7, X8, X9]:(p__d__disjointDecomposition3(X7,X8,X9)<=>p__d__disjoint(X8,X9))&![X7, X8, X9, X10]:(p__d__disjointDecomposition4(X7,X8,X9,X10)<=>((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X9,X10))))&![X7, X8, X9, X10, X11]:(p__d__disjointDecomposition5(X7,X8,X9,X10,X11)<=>(((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X10,X11))))&![X7, X8, X9, X10, X11, X12]:(p__d__disjointDecomposition6(X7,X8,X9,X10,X11,X12)<=>(((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X11,X12))))&![X7, X8, X9, X10, X11, X12, X13]:(p__d__disjointDecomposition7(X7,X8,X9,X10,X11,X12,X13)<=>((((((((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X8,X13))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X9,X13))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X10,X13))&p__d__disjoint(X11,X12))&p__d__disjoint(X11,X13))&p__d__disjoint(X12,X13)))), file('adimen.sumo.tstp', predefinitionsA24)).\\nfof(predefinitionsA8, axiom, ![X1, X2, X3]:((p__d__subclass(X1,X2)&p__d__subclass(X2,X3))=>p__d__subclass(X1,X3)), file('adimen.sumo.tstp', predefinitionsA8)).\\nfof(mergeA3615, axiom, p__d__subclass(c__ClosedTwoDimensionalFigure,c__TwoDimensionalFigure), file('adimen.sumo.tstp', mergeA3615)).\\nfof(mergeA331, axiom, p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition), file('adimen.sumo.tstp', mergeA331)).\\nfof(mergeA333, axiom, p__d__subclass(c__Attribute,c__Abstract), file('adimen.sumo.tstp', mergeA333)).\\nfof(mergeA330, axiom, p__d__subclass(c__Abstract,c__Entity), file('adimen.sumo.tstp', mergeA330)).\\nfof(mergeA351, axiom, p__d__subclass(c__InternalAttribute,c__Attribute), file('adimen.sumo.tstp', mergeA351)).\\nfof(mergeA3596, axiom, p__d__subclass(c__ShapeAttribute,c__StructureAttribute), file('adimen.sumo.tstp', mergeA3596)).\\nfof(mergeA3595, axiom, p__d__subclass(c__StructureAttribute,c__InternalAttribute), file('adimen.sumo.tstp', mergeA3595)).\\nfof(mergeA3607, axiom, p__d__subclass(c__GeometricFigure,c__ShapeAttribute), file('adimen.sumo.tstp', mergeA3607)).\\nfof(mergeA3611, axiom, p__d__subclass(c__TwoDimensionalFigure,c__GeometricFigure), file('adimen.sumo.tstp', mergeA3611)).\\nfof(mergeA575, axiom, p__d__subclass(c__Class,c__SetOrClass), file('adimen.sumo.tstp', mergeA575)).\\nfof(mergeA177, axiom, ![X7]:(p__d__instance(X7,c__Class)<=>p__d__subclass(X7,c__Entity)), file('adimen.sumo.tstp', mergeA177)).\\nfof(c_0_17, negated_conjecture, ~(?[X610]:~p__d__instance(X610,c__Polygon)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[whiteBoxTruthTest4282])])).\\nfof(c_0_18, plain, ![X7776, X7777, X7778]:(~p__d__instance(X7776,X7777)|~p__d__subclass(X7777,X7778)|p__d__instance(X7776,X7778)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA12])])).\\nfof(c_0_19, negated_conjecture, ![X14607]:p__d__instance(X14607,c__Polygon), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).\\ncnf(c_0_20, plain, (p__d__instance(X1,X3)|~p__d__instance(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).\\ncnf(c_0_21, plain, (p__d__subclass(c__Polygon,c__ClosedTwoDimensionalFigure)), inference(split_conjunct,[status(thm)],[miloA759])).\\ncnf(c_0_22, negated_conjecture, (p__d__instance(X1,c__Polygon)), inference(split_conjunct,[status(thm)],[c_0_19])).\\nfof(c_0_23, plain, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~p__d__instance(X6,X4)|~p__d__instance(X6,X5))), inference(fof_simplification,[status(thm)],[predefinitionsA15])).\\nfof(c_0_24, plain, ![X7870, X7871, X7872, X7873, X7874, X7875, X7876, X7877, X7878, X7879, X7880, X7881, X7882, X7883, X7884, X7885, X7886, X7887, X7888, X7889, X7890, X7891, X7892, X7893, X7894]:((((((~p__d__disjointDecomposition3(X7870,X7871,X7872)|p__d__disjoint(X7871,X7872))&(~p__d__disjoint(X7871,X7872)|p__d__disjointDecomposition3(X7870,X7871,X7872)))&((((p__d__disjoint(X7874,X7875)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))&(p__d__disjoint(X7874,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(p__d__disjoint(X7875,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(~p__d__disjoint(X7874,X7875)|~p__d__disjoint(X7874,X7876)|~p__d__disjoint(X7875,X7876)|p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))))&(((((((p__d__disjoint(X7878,X7879)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))&(p__d__disjoint(X7878,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7878,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7880,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(~p__d__disjoint(X7878,X7879)|~p__d__disjoint(X7878,X7880)|~p__d__disjoint(X7878,X7881)|~p__d__disjoint(X7879,X7880)|~p__d__disjoint(X7879,X7881)|~p__d__disjoint(X7880,X7881)|p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))))&(((((((((((p__d__disjoint(X7883,X7884)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))&(p__d__disjoint(X7883,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7886,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(~p__d__disjoint(X7883,X7884)|~p__d__disjoint(X7883,X7885)|~p__d__disjoint(X7883,X7886)|~p__d__disjoint(X7883,X7887)|~p__d__disjoint(X7884,X7885)|~p__d__disjoint(X7884,X7886)|~p__d__disjoint(X7884,X7887)|~p__d__disjoint(X7885,X7886)|~p__d__disjoint(X7885,X7887)|~p__d__disjoint(X7886,X7887)|p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))))&((((((((((((((((p__d__disjoint(X7889,X7890)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894))&(p__d__disjoint(X7889,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7893,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(~p__d__disjoint(X7889,X7890)|~p__d__disjoint(X7889,X7891)|~p__d__disjoint(X7889,X7892)|~p__d__disjoint(X7889,X7893)|~p__d__disjoint(X7889,X7894)|~p__d__disjoint(X7890,X7891)|~p__d__disjoint(X7890,X7892)|~p__d__disjoint(X7890,X7893)|~p__d__disjoint(X7890,X7894)|~p__d__disjoint(X7891,X7892)|~p__d__disjoint(X7891,X7893)|~p__d__disjoint(X7891,X7894)|~p__d__disjoint(X7892,X7893)|~p__d__disjoint(X7892,X7894)|~p__d__disjoint(X7893,X7894)|p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA24])])])])).\\nfof(c_0_25, plain, ![X7771, X7772, X7773]:(~p__d__subclass(X7771,X7772)|~p__d__subclass(X7772,X7773)|p__d__subclass(X7771,X7773)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA8])])).\\ncnf(c_0_26, plain, (p__d__subclass(c__ClosedTwoDimensionalFigure,c__TwoDimensionalFigure)), inference(split_conjunct,[status(thm)],[mergeA3615])).\\ncnf(c_0_27, plain, (p__d__instance(X1,c__ClosedTwoDimensionalFigure)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])])).\\nfof(c_0_28, plain, ![X7779, X7780, X7781, X7782, X7783]:((~p__d__disjoint(X7779,X7780)|(~p__d__instance(X7781,X7779)|~p__d__instance(X7781,X7780)))&((p__d__instance(esk1_2(X7782,X7783),X7782)|p__d__disjoint(X7782,X7783))&(p__d__instance(esk1_2(X7782,X7783),X7783)|p__d__disjoint(X7782,X7783)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])])])).\\ncnf(c_0_29, plain, (p__d__disjoint(X1,X2)|~p__d__disjointDecomposition6(X3,X4,X1,X2,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_24])).\\ncnf(c_0_30, plain, (p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition)), inference(split_conjunct,[status(thm)],[mergeA331])).\\ncnf(c_0_31, plain, (p__d__subclass(X1,X3)|~p__d__subclass(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).\\ncnf(c_0_32, plain, (p__d__subclass(c__Attribute,c__Abstract)), inference(split_conjunct,[status(thm)],[mergeA333])).\\ncnf(c_0_33, plain, (p__d__subclass(c__Abstract,c__Entity)), inference(split_conjunct,[status(thm)],[mergeA330])).\\ncnf(c_0_34, plain, (p__d__subclass(c__InternalAttribute,c__Attribute)), inference(split_conjunct,[status(thm)],[mergeA351])).\\ncnf(c_0_35, plain, (p__d__subclass(c__ShapeAttribute,c__StructureAttribute)), inference(split_conjunct,[status(thm)],[mergeA3596])).\\ncnf(c_0_36, plain, (p__d__subclass(c__StructureAttribute,c__InternalAttribute)), inference(split_conjunct,[status(thm)],[mergeA3595])).\\ncnf(c_0_37, plain, (p__d__subclass(c__GeometricFigure,c__ShapeAttribute)), inference(split_conjunct,[status(thm)],[mergeA3607])).\\ncnf(c_0_38, plain, (p__d__subclass(c__TwoDimensionalFigure,c__GeometricFigure)), inference(split_conjunct,[status(thm)],[mergeA3611])).\\ncnf(c_0_39, plain, (p__d__subclass(c__Class,c__SetOrClass)), inference(split_conjunct,[status(thm)],[mergeA575])).\\ncnf(c_0_40, plain, (p__d__instance(X1,c__TwoDimensionalFigure)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_26]), c_0_27])])).\\ncnf(c_0_41, plain, (~p__d__disjoint(X1,X2)|~p__d__instance(X3,X1)|~p__d__instance(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).\\ncnf(c_0_42, plain, (p__d__disjoint(c__Attribute,c__SetOrClass)), inference(spm,[status(thm)],[c_0_29, c_0_30])).\\nfof(c_0_43, plain, ![X8007]:((~p__d__instance(X8007,c__Class)|p__d__subclass(X8007,c__Entity))&(~p__d__subclass(X8007,c__Entity)|p__d__instance(X8007,c__Class))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mergeA177])])).\\ncnf(c_0_44, plain, (p__d__subclass(X1,c__Abstract)|~p__d__subclass(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_31, c_0_32])).\\ncnf(c_0_45, plain, (p__d__subclass(X1,c__Entity)|~p__d__subclass(X1,c__Abstract)), inference(spm,[status(thm)],[c_0_31, c_0_33])).\\ncnf(c_0_46, plain, (p__d__subclass(X1,c__Attribute)|~p__d__subclass(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_31, c_0_34])).\\ncnf(c_0_47, plain, (p__d__subclass(X1,c__TwoDimensionalFigure)|~p__d__subclass(X1,c__ClosedTwoDimensionalFigure)), inference(spm,[status(thm)],[c_0_31, c_0_26])).\\ncnf(c_0_48, plain, (p__d__subclass(X1,c__StructureAttribute)|~p__d__subclass(X1,c__ShapeAttribute)), inference(spm,[status(thm)],[c_0_31, c_0_35])).\\ncnf(c_0_49, plain, (p__d__subclass(X1,c__InternalAttribute)|~p__d__subclass(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_31, c_0_36])).\\ncnf(c_0_50, plain, (p__d__subclass(X1,c__ShapeAttribute)|~p__d__subclass(X1,c__GeometricFigure)), inference(spm,[status(thm)],[c_0_31, c_0_37])).\\ncnf(c_0_51, plain, (p__d__subclass(X1,c__GeometricFigure)|~p__d__subclass(X1,c__TwoDimensionalFigure)), inference(spm,[status(thm)],[c_0_31, c_0_38])).\\ncnf(c_0_52, plain, (p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Class)), inference(spm,[status(thm)],[c_0_20, c_0_39])).\\ncnf(c_0_53, plain, (p__d__instance(X1,c__Attribute)|~p__d__instance(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_20, c_0_34])).\\ncnf(c_0_54, plain, (p__d__instance(X1,c__StructureAttribute)|~p__d__instance(X1,c__ShapeAttribute)), inference(spm,[status(thm)],[c_0_20, c_0_35])).\\ncnf(c_0_55, plain, (p__d__instance(X1,c__InternalAttribute)|~p__d__instance(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_20, c_0_36])).\\ncnf(c_0_56, plain, (p__d__instance(X1,c__ShapeAttribute)|~p__d__instance(X1,c__GeometricFigure)), inference(spm,[status(thm)],[c_0_20, c_0_37])).\\ncnf(c_0_57, plain, (p__d__instance(X1,c__GeometricFigure)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_38]), c_0_40])])).\\ncnf(c_0_58, plain, (~p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_41, c_0_42])).\\ncnf(c_0_59, plain, (p__d__instance(X1,c__Class)|~p__d__subclass(X1,c__Entity)), inference(split_conjunct,[status(thm)],[c_0_43])).\\ncnf(c_0_60, plain, ($false), inference(cdclpropres,[status(thm)],[c_0_44, c_0_45, c_0_46, c_0_47, c_0_48, c_0_49, c_0_50, c_0_51, c_0_52, c_0_53, c_0_54, c_0_55, c_0_56, c_0_57, c_0_58, c_0_59, c_0_21]), ['proof']).\\n# SZS output end CNFRefutation\\n# Training examples: 0 positive, 0 negative\\n\""
  ],
  "./goals/whitebox/whiteBoxTruthTest4309.adimen.tstp": [
    "ProverResult.PROOF_FOUND",
    "b\"# No SInE strategy applied\\n# Auto-Mode selected heuristic C07_19_nc_SAT001_MinMin_rr\\n# and selection function SelectMaxLComplexAvoidPosPred.\\n#\\n# Presaturation interreduction done\\n# SatCheck found unsatisfiable ground set\\n\\n# Proof found!\\n# SZS status Theorem\\n# SZS output start CNFRefutation\\nfof(whiteBoxTruthTest4309, conjecture, ?[X612, X175]:![X213]:(~(p__d__instance(X213,c__OneDimensionalFigure))|~(p__sideOfFigure(X213,X612))|p__lineMeasure(X213,X175)), file('./goals/whitebox/whiteBoxTruthTest4309.adimen.tstp', whiteBoxTruthTest4309)).\\nfof(predefinitionsA15, axiom, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~(p__d__instance(X6,X4))|~(p__d__instance(X6,X5)))), file('adimen.sumo.tstp', predefinitionsA15)).\\nfof(predefinitionsA24, axiom, ((((![X7, X8, X9]:(p__d__disjointDecomposition3(X7,X8,X9)<=>p__d__disjoint(X8,X9))&![X7, X8, X9, X10]:(p__d__disjointDecomposition4(X7,X8,X9,X10)<=>((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X9,X10))))&![X7, X8, X9, X10, X11]:(p__d__disjointDecomposition5(X7,X8,X9,X10,X11)<=>(((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X10,X11))))&![X7, X8, X9, X10, X11, X12]:(p__d__disjointDecomposition6(X7,X8,X9,X10,X11,X12)<=>(((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X11,X12))))&![X7, X8, X9, X10, X11, X12, X13]:(p__d__disjointDecomposition7(X7,X8,X9,X10,X11,X12,X13)<=>((((((((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X8,X13))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X9,X13))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X10,X13))&p__d__disjoint(X11,X12))&p__d__disjoint(X11,X13))&p__d__disjoint(X12,X13)))), file('adimen.sumo.tstp', predefinitionsA24)).\\nfof(typeA419, axiom, ![X1038, X1039]:(p__sideOfFigure(X1038,X1039)=>(p__d__instance(X1039,c__GeometricFigure)&p__d__instance(X1038,c__OneDimensionalFigure))), file('adimen.sumo.tstp', typeA419)).\\nfof(predefinitionsA12, axiom, ![X1, X2, X3]:((p__d__instance(X1,X2)&p__d__subclass(X2,X3))=>p__d__instance(X1,X3)), file('adimen.sumo.tstp', predefinitionsA12)).\\nfof(predefinitionsA8, axiom, ![X1, X2, X3]:((p__d__subclass(X1,X2)&p__d__subclass(X2,X3))=>p__d__subclass(X1,X3)), file('adimen.sumo.tstp', predefinitionsA8)).\\nfof(mergeA331, axiom, p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition), file('adimen.sumo.tstp', mergeA331)).\\nfof(mergeA575, axiom, p__d__subclass(c__Class,c__SetOrClass), file('adimen.sumo.tstp', mergeA575)).\\nfof(mergeA351, axiom, p__d__subclass(c__InternalAttribute,c__Attribute), file('adimen.sumo.tstp', mergeA351)).\\nfof(mergeA3596, axiom, p__d__subclass(c__ShapeAttribute,c__StructureAttribute), file('adimen.sumo.tstp', mergeA3596)).\\nfof(mergeA333, axiom, p__d__subclass(c__Attribute,c__Abstract), file('adimen.sumo.tstp', mergeA333)).\\nfof(mergeA330, axiom, p__d__subclass(c__Abstract,c__Entity), file('adimen.sumo.tstp', mergeA330)).\\nfof(mergeA3595, axiom, p__d__subclass(c__StructureAttribute,c__InternalAttribute), file('adimen.sumo.tstp', mergeA3595)).\\nfof(mergeA3607, axiom, p__d__subclass(c__GeometricFigure,c__ShapeAttribute), file('adimen.sumo.tstp', mergeA3607)).\\nfof(mergeA177, axiom, ![X7]:(p__d__instance(X7,c__Class)<=>p__d__subclass(X7,c__Entity)), file('adimen.sumo.tstp', mergeA177)).\\nfof(mergeA3610, axiom, p__d__subclass(c__OneDimensionalFigure,c__GeometricFigure), file('adimen.sumo.tstp', mergeA3610)).\\nfof(c_0_16, negated_conjecture, ~(?[X612, X175]:![X213]:(~p__d__instance(X213,c__OneDimensionalFigure)|~p__sideOfFigure(X213,X612)|p__lineMeasure(X213,X175))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[whiteBoxTruthTest4309])])).\\nfof(c_0_17, plain, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~p__d__instance(X6,X4)|~p__d__instance(X6,X5))), inference(fof_simplification,[status(thm)],[predefinitionsA15])).\\nfof(c_0_18, plain, ![X7872, X7873, X7874, X7875, X7876, X7877, X7878, X7879, X7880, X7881, X7882, X7883, X7884, X7885, X7886, X7887, X7888, X7889, X7890, X7891, X7892, X7893, X7894, X7895, X7896]:((((((~p__d__disjointDecomposition3(X7872,X7873,X7874)|p__d__disjoint(X7873,X7874))&(~p__d__disjoint(X7873,X7874)|p__d__disjointDecomposition3(X7872,X7873,X7874)))&((((p__d__disjoint(X7876,X7877)|~p__d__disjointDecomposition4(X7875,X7876,X7877,X7878))&(p__d__disjoint(X7876,X7878)|~p__d__disjointDecomposition4(X7875,X7876,X7877,X7878)))&(p__d__disjoint(X7877,X7878)|~p__d__disjointDecomposition4(X7875,X7876,X7877,X7878)))&(~p__d__disjoint(X7876,X7877)|~p__d__disjoint(X7876,X7878)|~p__d__disjoint(X7877,X7878)|p__d__disjointDecomposition4(X7875,X7876,X7877,X7878))))&(((((((p__d__disjoint(X7880,X7881)|~p__d__disjointDecomposition5(X7879,X7880,X7881,X7882,X7883))&(p__d__disjoint(X7880,X7882)|~p__d__disjointDecomposition5(X7879,X7880,X7881,X7882,X7883)))&(p__d__disjoint(X7880,X7883)|~p__d__disjointDecomposition5(X7879,X7880,X7881,X7882,X7883)))&(p__d__disjoint(X7881,X7882)|~p__d__disjointDecomposition5(X7879,X7880,X7881,X7882,X7883)))&(p__d__disjoint(X7881,X7883)|~p__d__disjointDecomposition5(X7879,X7880,X7881,X7882,X7883)))&(p__d__disjoint(X7882,X7883)|~p__d__disjointDecomposition5(X7879,X7880,X7881,X7882,X7883)))&(~p__d__disjoint(X7880,X7881)|~p__d__disjoint(X7880,X7882)|~p__d__disjoint(X7880,X7883)|~p__d__disjoint(X7881,X7882)|~p__d__disjoint(X7881,X7883)|~p__d__disjoint(X7882,X7883)|p__d__disjointDecomposition5(X7879,X7880,X7881,X7882,X7883))))&(((((((((((p__d__disjoint(X7885,X7886)|~p__d__disjointDecomposition6(X7884,X7885,X7886,X7887,X7888,X7889))&(p__d__disjoint(X7885,X7887)|~p__d__disjointDecomposition6(X7884,X7885,X7886,X7887,X7888,X7889)))&(p__d__disjoint(X7885,X7888)|~p__d__disjointDecomposition6(X7884,X7885,X7886,X7887,X7888,X7889)))&(p__d__disjoint(X7885,X7889)|~p__d__disjointDecomposition6(X7884,X7885,X7886,X7887,X7888,X7889)))&(p__d__disjoint(X7886,X7887)|~p__d__disjointDecomposition6(X7884,X7885,X7886,X7887,X7888,X7889)))&(p__d__disjoint(X7886,X7888)|~p__d__disjointDecomposition6(X7884,X7885,X7886,X7887,X7888,X7889)))&(p__d__disjoint(X7886,X7889)|~p__d__disjointDecomposition6(X7884,X7885,X7886,X7887,X7888,X7889)))&(p__d__disjoint(X7887,X7888)|~p__d__disjointDecomposition6(X7884,X7885,X7886,X7887,X7888,X7889)))&(p__d__disjoint(X7887,X7889)|~p__d__disjointDecomposition6(X7884,X7885,X7886,X7887,X7888,X7889)))&(p__d__disjoint(X7888,X7889)|~p__d__disjointDecomposition6(X7884,X7885,X7886,X7887,X7888,X7889)))&(~p__d__disjoint(X7885,X7886)|~p__d__disjoint(X7885,X7887)|~p__d__disjoint(X7885,X7888)|~p__d__disjoint(X7885,X7889)|~p__d__disjoint(X7886,X7887)|~p__d__disjoint(X7886,X7888)|~p__d__disjoint(X7886,X7889)|~p__d__disjoint(X7887,X7888)|~p__d__disjoint(X7887,X7889)|~p__d__disjoint(X7888,X7889)|p__d__disjointDecomposition6(X7884,X7885,X7886,X7887,X7888,X7889))))&((((((((((((((((p__d__disjoint(X7891,X7892)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896))&(p__d__disjoint(X7891,X7893)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7891,X7894)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7891,X7895)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7891,X7896)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7892,X7893)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7892,X7894)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7892,X7895)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7892,X7896)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7893,X7894)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7893,X7895)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7893,X7896)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7894,X7895)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7894,X7896)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(p__d__disjoint(X7895,X7896)|~p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))&(~p__d__disjoint(X7891,X7892)|~p__d__disjoint(X7891,X7893)|~p__d__disjoint(X7891,X7894)|~p__d__disjoint(X7891,X7895)|~p__d__disjoint(X7891,X7896)|~p__d__disjoint(X7892,X7893)|~p__d__disjoint(X7892,X7894)|~p__d__disjoint(X7892,X7895)|~p__d__disjoint(X7892,X7896)|~p__d__disjoint(X7893,X7894)|~p__d__disjoint(X7893,X7895)|~p__d__disjoint(X7893,X7896)|~p__d__disjoint(X7894,X7895)|~p__d__disjoint(X7894,X7896)|~p__d__disjoint(X7895,X7896)|p__d__disjointDecomposition7(X7890,X7891,X7892,X7893,X7894,X7895,X7896)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA24])])])])).\\nfof(c_0_19, plain, ![X14173, X14174]:((p__d__instance(X14174,c__GeometricFigure)|~p__sideOfFigure(X14173,X14174))&(p__d__instance(X14173,c__OneDimensionalFigure)|~p__sideOfFigure(X14173,X14174))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[typeA419])])])).\\nfof(c_0_20, negated_conjecture, ![X14609, X14610]:((p__d__instance(esk1216_2(X14609,X14610),c__OneDimensionalFigure)&p__sideOfFigure(esk1216_2(X14609,X14610),X14609))&~p__lineMeasure(esk1216_2(X14609,X14610),X14610)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).\\nfof(c_0_21, plain, ![X7778, X7779, X7780]:(~p__d__instance(X7778,X7779)|~p__d__subclass(X7779,X7780)|p__d__instance(X7778,X7780)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA12])])).\\nfof(c_0_22, plain, ![X7773, X7774, X7775]:(~p__d__subclass(X7773,X7774)|~p__d__subclass(X7774,X7775)|p__d__subclass(X7773,X7775)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA8])])).\\nfof(c_0_23, plain, ![X7781, X7782, X7783, X7784, X7785]:((~p__d__disjoint(X7781,X7782)|(~p__d__instance(X7783,X7781)|~p__d__instance(X7783,X7782)))&((p__d__instance(esk1_2(X7784,X7785),X7784)|p__d__disjoint(X7784,X7785))&(p__d__instance(esk1_2(X7784,X7785),X7785)|p__d__disjoint(X7784,X7785)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])])])).\\ncnf(c_0_24, plain, (p__d__disjoint(X1,X2)|~p__d__disjointDecomposition6(X3,X4,X1,X2,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_18])).\\ncnf(c_0_25, plain, (p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition)), inference(split_conjunct,[status(thm)],[mergeA331])).\\ncnf(c_0_26, plain, (p__d__instance(X1,c__GeometricFigure)|~p__sideOfFigure(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).\\ncnf(c_0_27, negated_conjecture, (p__sideOfFigure(esk1216_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_20])).\\ncnf(c_0_28, plain, (p__d__instance(X1,X3)|~p__d__instance(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).\\ncnf(c_0_29, plain, (p__d__subclass(c__Class,c__SetOrClass)), inference(split_conjunct,[status(thm)],[mergeA575])).\\ncnf(c_0_30, plain, (p__d__subclass(c__InternalAttribute,c__Attribute)), inference(split_conjunct,[status(thm)],[mergeA351])).\\ncnf(c_0_31, plain, (p__d__subclass(c__ShapeAttribute,c__StructureAttribute)), inference(split_conjunct,[status(thm)],[mergeA3596])).\\ncnf(c_0_32, plain, (p__d__subclass(X1,X3)|~p__d__subclass(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).\\ncnf(c_0_33, plain, (p__d__subclass(c__Attribute,c__Abstract)), inference(split_conjunct,[status(thm)],[mergeA333])).\\ncnf(c_0_34, plain, (p__d__subclass(c__Abstract,c__Entity)), inference(split_conjunct,[status(thm)],[mergeA330])).\\ncnf(c_0_35, plain, (p__d__subclass(c__StructureAttribute,c__InternalAttribute)), inference(split_conjunct,[status(thm)],[mergeA3595])).\\ncnf(c_0_36, plain, (p__d__subclass(c__GeometricFigure,c__ShapeAttribute)), inference(split_conjunct,[status(thm)],[mergeA3607])).\\ncnf(c_0_37, plain, (~p__d__disjoint(X1,X2)|~p__d__instance(X3,X1)|~p__d__instance(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).\\ncnf(c_0_38, plain, (p__d__disjoint(c__Attribute,c__SetOrClass)), inference(spm,[status(thm)],[c_0_24, c_0_25])).\\nfof(c_0_39, plain, ![X8009]:((~p__d__instance(X8009,c__Class)|p__d__subclass(X8009,c__Entity))&(~p__d__subclass(X8009,c__Entity)|p__d__instance(X8009,c__Class))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mergeA177])])).\\ncnf(c_0_40, negated_conjecture, (p__d__instance(X1,c__GeometricFigure)), inference(spm,[status(thm)],[c_0_26, c_0_27])).\\ncnf(c_0_41, plain, (p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Class)), inference(spm,[status(thm)],[c_0_28, c_0_29])).\\ncnf(c_0_42, plain, (p__d__instance(X1,c__Attribute)|~p__d__instance(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_28, c_0_30])).\\ncnf(c_0_43, plain, (p__d__instance(X1,c__StructureAttribute)|~p__d__instance(X1,c__ShapeAttribute)), inference(spm,[status(thm)],[c_0_28, c_0_31])).\\ncnf(c_0_44, plain, (p__d__subclass(X1,c__Abstract)|~p__d__subclass(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_32, c_0_33])).\\ncnf(c_0_45, plain, (p__d__subclass(X1,c__Entity)|~p__d__subclass(X1,c__Abstract)), inference(spm,[status(thm)],[c_0_32, c_0_34])).\\ncnf(c_0_46, plain, (p__d__subclass(X1,c__Attribute)|~p__d__subclass(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_32, c_0_30])).\\ncnf(c_0_47, plain, (p__d__subclass(X1,c__StructureAttribute)|~p__d__subclass(X1,c__ShapeAttribute)), inference(spm,[status(thm)],[c_0_32, c_0_31])).\\ncnf(c_0_48, plain, (p__d__subclass(X1,c__InternalAttribute)|~p__d__subclass(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_32, c_0_35])).\\ncnf(c_0_49, plain, (p__d__subclass(X1,c__ShapeAttribute)|~p__d__subclass(X1,c__GeometricFigure)), inference(spm,[status(thm)],[c_0_32, c_0_36])).\\ncnf(c_0_50, plain, (p__d__instance(X1,c__InternalAttribute)|~p__d__instance(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_28, c_0_35])).\\ncnf(c_0_51, plain, (~p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_37, c_0_38])).\\ncnf(c_0_52, plain, (p__d__instance(X1,c__Class)|~p__d__subclass(X1,c__Entity)), inference(split_conjunct,[status(thm)],[c_0_39])).\\ncnf(c_0_53, plain, (p__d__instance(X1,c__ShapeAttribute)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_36]), c_0_40])])).\\ncnf(c_0_54, plain, (p__d__subclass(c__OneDimensionalFigure,c__GeometricFigure)), inference(split_conjunct,[status(thm)],[mergeA3610])).\\ncnf(c_0_55, plain, ($false), inference(cdclpropres,[status(thm)],[c_0_41, c_0_42, c_0_43, c_0_44, c_0_45, c_0_46, c_0_47, c_0_48, c_0_49, c_0_50, c_0_51, c_0_52, c_0_53, c_0_54]), ['proof']).\\n# SZS output end CNFRefutation\\n# Training examples: 0 positive, 0 negative\\n\""
  ],
  "./goals/whitebox/whiteBoxTruthTest7727.adimen.tstp": [
    "ProverResult.PROOF_FOUND",
    "b\"# No SInE strategy applied\\n# Auto-Mode selected heuristic SAT001_MinMin_NN\\n# and selection function SelectMaxLComplexAvoidPosPred.\\n#\\n# Presaturation interreduction done\\n# SatCheck found unsatisfiable ground set\\n\\n# Proof found!\\n# SZS status Theorem\\n# SZS output start CNFRefutation\\nfof(whiteBoxTruthTest7727, conjecture, ?[X1017]:~(p__d__instance(X1017,c__HemisphereFigure)), file('./goals/whitebox/whiteBoxTruthTest7727.adimen.tstp', whiteBoxTruthTest7727)).\\nfof(predefinitionsA12, axiom, ![X1, X2, X3]:((p__d__instance(X1,X2)&p__d__subclass(X2,X3))=>p__d__instance(X1,X3)), file('adimen.sumo.tstp', predefinitionsA12)).\\nfof(miloA4156, axiom, p__d__subclass(c__HemisphereFigure,c__ThreeDimensionalFigure), file('adimen.sumo.tstp', miloA4156)).\\nfof(predefinitionsA15, axiom, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~(p__d__instance(X6,X4))|~(p__d__instance(X6,X5)))), file('adimen.sumo.tstp', predefinitionsA15)).\\nfof(predefinitionsA24, axiom, ((((![X7, X8, X9]:(p__d__disjointDecomposition3(X7,X8,X9)<=>p__d__disjoint(X8,X9))&![X7, X8, X9, X10]:(p__d__disjointDecomposition4(X7,X8,X9,X10)<=>((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X9,X10))))&![X7, X8, X9, X10, X11]:(p__d__disjointDecomposition5(X7,X8,X9,X10,X11)<=>(((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X10,X11))))&![X7, X8, X9, X10, X11, X12]:(p__d__disjointDecomposition6(X7,X8,X9,X10,X11,X12)<=>(((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X11,X12))))&![X7, X8, X9, X10, X11, X12, X13]:(p__d__disjointDecomposition7(X7,X8,X9,X10,X11,X12,X13)<=>((((((((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X8,X13))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X9,X13))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X10,X13))&p__d__disjoint(X11,X12))&p__d__disjoint(X11,X13))&p__d__disjoint(X12,X13)))), file('adimen.sumo.tstp', predefinitionsA24)).\\nfof(predefinitionsA8, axiom, ![X1, X2, X3]:((p__d__subclass(X1,X2)&p__d__subclass(X2,X3))=>p__d__subclass(X1,X3)), file('adimen.sumo.tstp', predefinitionsA8)).\\nfof(mergeA3616, axiom, p__d__subclass(c__ThreeDimensionalFigure,c__GeometricFigure), file('adimen.sumo.tstp', mergeA3616)).\\nfof(mergeA331, axiom, p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition), file('adimen.sumo.tstp', mergeA331)).\\nfof(mergeA333, axiom, p__d__subclass(c__Attribute,c__Abstract), file('adimen.sumo.tstp', mergeA333)).\\nfof(mergeA330, axiom, p__d__subclass(c__Abstract,c__Entity), file('adimen.sumo.tstp', mergeA330)).\\nfof(mergeA351, axiom, p__d__subclass(c__InternalAttribute,c__Attribute), file('adimen.sumo.tstp', mergeA351)).\\nfof(mergeA3596, axiom, p__d__subclass(c__ShapeAttribute,c__StructureAttribute), file('adimen.sumo.tstp', mergeA3596)).\\nfof(mergeA3595, axiom, p__d__subclass(c__StructureAttribute,c__InternalAttribute), file('adimen.sumo.tstp', mergeA3595)).\\nfof(mergeA3607, axiom, p__d__subclass(c__GeometricFigure,c__ShapeAttribute), file('adimen.sumo.tstp', mergeA3607)).\\nfof(mergeA575, axiom, p__d__subclass(c__Class,c__SetOrClass), file('adimen.sumo.tstp', mergeA575)).\\nfof(mergeA177, axiom, ![X7]:(p__d__instance(X7,c__Class)<=>p__d__subclass(X7,c__Entity)), file('adimen.sumo.tstp', mergeA177)).\\nfof(c_0_16, negated_conjecture, ~(?[X1017]:~p__d__instance(X1017,c__HemisphereFigure)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[whiteBoxTruthTest7727])])).\\nfof(c_0_17, plain, ![X7776, X7777, X7778]:(~p__d__instance(X7776,X7777)|~p__d__subclass(X7777,X7778)|p__d__instance(X7776,X7778)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA12])])).\\nfof(c_0_18, negated_conjecture, ![X14607]:p__d__instance(X14607,c__HemisphereFigure), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).\\ncnf(c_0_19, plain, (p__d__instance(X1,X3)|~p__d__instance(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).\\ncnf(c_0_20, plain, (p__d__subclass(c__HemisphereFigure,c__ThreeDimensionalFigure)), inference(split_conjunct,[status(thm)],[miloA4156])).\\ncnf(c_0_21, negated_conjecture, (p__d__instance(X1,c__HemisphereFigure)), inference(split_conjunct,[status(thm)],[c_0_18])).\\nfof(c_0_22, plain, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~p__d__instance(X6,X4)|~p__d__instance(X6,X5))), inference(fof_simplification,[status(thm)],[predefinitionsA15])).\\nfof(c_0_23, plain, ![X7870, X7871, X7872, X7873, X7874, X7875, X7876, X7877, X7878, X7879, X7880, X7881, X7882, X7883, X7884, X7885, X7886, X7887, X7888, X7889, X7890, X7891, X7892, X7893, X7894]:((((((~p__d__disjointDecomposition3(X7870,X7871,X7872)|p__d__disjoint(X7871,X7872))&(~p__d__disjoint(X7871,X7872)|p__d__disjointDecomposition3(X7870,X7871,X7872)))&((((p__d__disjoint(X7874,X7875)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))&(p__d__disjoint(X7874,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(p__d__disjoint(X7875,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(~p__d__disjoint(X7874,X7875)|~p__d__disjoint(X7874,X7876)|~p__d__disjoint(X7875,X7876)|p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))))&(((((((p__d__disjoint(X7878,X7879)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))&(p__d__disjoint(X7878,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7878,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7880,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(~p__d__disjoint(X7878,X7879)|~p__d__disjoint(X7878,X7880)|~p__d__disjoint(X7878,X7881)|~p__d__disjoint(X7879,X7880)|~p__d__disjoint(X7879,X7881)|~p__d__disjoint(X7880,X7881)|p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))))&(((((((((((p__d__disjoint(X7883,X7884)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))&(p__d__disjoint(X7883,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7886,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(~p__d__disjoint(X7883,X7884)|~p__d__disjoint(X7883,X7885)|~p__d__disjoint(X7883,X7886)|~p__d__disjoint(X7883,X7887)|~p__d__disjoint(X7884,X7885)|~p__d__disjoint(X7884,X7886)|~p__d__disjoint(X7884,X7887)|~p__d__disjoint(X7885,X7886)|~p__d__disjoint(X7885,X7887)|~p__d__disjoint(X7886,X7887)|p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))))&((((((((((((((((p__d__disjoint(X7889,X7890)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894))&(p__d__disjoint(X7889,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7893,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(~p__d__disjoint(X7889,X7890)|~p__d__disjoint(X7889,X7891)|~p__d__disjoint(X7889,X7892)|~p__d__disjoint(X7889,X7893)|~p__d__disjoint(X7889,X7894)|~p__d__disjoint(X7890,X7891)|~p__d__disjoint(X7890,X7892)|~p__d__disjoint(X7890,X7893)|~p__d__disjoint(X7890,X7894)|~p__d__disjoint(X7891,X7892)|~p__d__disjoint(X7891,X7893)|~p__d__disjoint(X7891,X7894)|~p__d__disjoint(X7892,X7893)|~p__d__disjoint(X7892,X7894)|~p__d__disjoint(X7893,X7894)|p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA24])])])])).\\nfof(c_0_24, plain, ![X7771, X7772, X7773]:(~p__d__subclass(X7771,X7772)|~p__d__subclass(X7772,X7773)|p__d__subclass(X7771,X7773)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA8])])).\\ncnf(c_0_25, plain, (p__d__subclass(c__ThreeDimensionalFigure,c__GeometricFigure)), inference(split_conjunct,[status(thm)],[mergeA3616])).\\ncnf(c_0_26, plain, (p__d__instance(X1,c__ThreeDimensionalFigure)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).\\nfof(c_0_27, plain, ![X7779, X7780, X7781, X7782, X7783]:((~p__d__disjoint(X7779,X7780)|(~p__d__instance(X7781,X7779)|~p__d__instance(X7781,X7780)))&((p__d__instance(esk1_2(X7782,X7783),X7782)|p__d__disjoint(X7782,X7783))&(p__d__instance(esk1_2(X7782,X7783),X7783)|p__d__disjoint(X7782,X7783)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])])])).\\ncnf(c_0_28, plain, (p__d__disjoint(X1,X2)|~p__d__disjointDecomposition6(X3,X4,X1,X2,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_23])).\\ncnf(c_0_29, plain, (p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition)), inference(split_conjunct,[status(thm)],[mergeA331])).\\ncnf(c_0_30, plain, (p__d__subclass(X1,X3)|~p__d__subclass(X1,X2)|~p__d__subclass(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_24])).\\ncnf(c_0_31, plain, (p__d__subclass(c__Attribute,c__Abstract)), inference(split_conjunct,[status(thm)],[mergeA333])).\\ncnf(c_0_32, plain, (p__d__subclass(c__Abstract,c__Entity)), inference(split_conjunct,[status(thm)],[mergeA330])).\\ncnf(c_0_33, plain, (p__d__subclass(c__InternalAttribute,c__Attribute)), inference(split_conjunct,[status(thm)],[mergeA351])).\\ncnf(c_0_34, plain, (p__d__subclass(c__ShapeAttribute,c__StructureAttribute)), inference(split_conjunct,[status(thm)],[mergeA3596])).\\ncnf(c_0_35, plain, (p__d__subclass(c__StructureAttribute,c__InternalAttribute)), inference(split_conjunct,[status(thm)],[mergeA3595])).\\ncnf(c_0_36, plain, (p__d__subclass(c__GeometricFigure,c__ShapeAttribute)), inference(split_conjunct,[status(thm)],[mergeA3607])).\\ncnf(c_0_37, plain, (p__d__subclass(c__Class,c__SetOrClass)), inference(split_conjunct,[status(thm)],[mergeA575])).\\ncnf(c_0_38, plain, (p__d__instance(X1,c__GeometricFigure)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_25]), c_0_26])])).\\ncnf(c_0_39, plain, (~p__d__disjoint(X1,X2)|~p__d__instance(X3,X1)|~p__d__instance(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).\\ncnf(c_0_40, plain, (p__d__disjoint(c__Attribute,c__SetOrClass)), inference(spm,[status(thm)],[c_0_28, c_0_29])).\\nfof(c_0_41, plain, ![X8007]:((~p__d__instance(X8007,c__Class)|p__d__subclass(X8007,c__Entity))&(~p__d__subclass(X8007,c__Entity)|p__d__instance(X8007,c__Class))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mergeA177])])).\\ncnf(c_0_42, plain, (p__d__subclass(X1,c__Abstract)|~p__d__subclass(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_30, c_0_31])).\\ncnf(c_0_43, plain, (p__d__subclass(X1,c__Entity)|~p__d__subclass(X1,c__Abstract)), inference(spm,[status(thm)],[c_0_30, c_0_32])).\\ncnf(c_0_44, plain, (p__d__subclass(X1,c__Attribute)|~p__d__subclass(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_30, c_0_33])).\\ncnf(c_0_45, plain, (p__d__subclass(X1,c__StructureAttribute)|~p__d__subclass(X1,c__ShapeAttribute)), inference(spm,[status(thm)],[c_0_30, c_0_34])).\\ncnf(c_0_46, plain, (p__d__subclass(X1,c__InternalAttribute)|~p__d__subclass(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_30, c_0_35])).\\ncnf(c_0_47, plain, (p__d__subclass(X1,c__ShapeAttribute)|~p__d__subclass(X1,c__GeometricFigure)), inference(spm,[status(thm)],[c_0_30, c_0_36])).\\ncnf(c_0_48, plain, (p__d__subclass(X1,c__GeometricFigure)|~p__d__subclass(X1,c__ThreeDimensionalFigure)), inference(spm,[status(thm)],[c_0_30, c_0_25])).\\ncnf(c_0_49, plain, (p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Class)), inference(spm,[status(thm)],[c_0_19, c_0_37])).\\ncnf(c_0_50, plain, (p__d__instance(X1,c__Attribute)|~p__d__instance(X1,c__InternalAttribute)), inference(spm,[status(thm)],[c_0_19, c_0_33])).\\ncnf(c_0_51, plain, (p__d__instance(X1,c__StructureAttribute)|~p__d__instance(X1,c__ShapeAttribute)), inference(spm,[status(thm)],[c_0_19, c_0_34])).\\ncnf(c_0_52, plain, (p__d__instance(X1,c__InternalAttribute)|~p__d__instance(X1,c__StructureAttribute)), inference(spm,[status(thm)],[c_0_19, c_0_35])).\\ncnf(c_0_53, plain, (p__d__instance(X1,c__ShapeAttribute)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_36]), c_0_38])])).\\ncnf(c_0_54, plain, (~p__d__instance(X1,c__SetOrClass)|~p__d__instance(X1,c__Attribute)), inference(spm,[status(thm)],[c_0_39, c_0_40])).\\ncnf(c_0_55, plain, (p__d__instance(X1,c__Class)|~p__d__subclass(X1,c__Entity)), inference(split_conjunct,[status(thm)],[c_0_41])).\\ncnf(c_0_56, plain, ($false), inference(cdclpropres,[status(thm)],[c_0_42, c_0_43, c_0_44, c_0_45, c_0_46, c_0_47, c_0_48, c_0_49, c_0_50, c_0_51, c_0_52, c_0_53, c_0_54, c_0_55, c_0_20]), ['proof']).\\n# SZS output end CNFRefutation\\n# Training examples: 0 positive, 0 negative\\n\""
  ],
  "./goals/whitebox/whiteBoxTruthTest0005.adimen.tstp": [
    "ProverResult.PROOF_FOUND",
    "b\"# No SInE strategy applied\\n# Auto-Mode selected heuristic SAT001_MinMin_NN\\n# and selection function SelectMaxLComplexAvoidPosPred.\\n#\\n# Presaturation interreduction done\\n\\n# Proof found!\\n# SZS status Theorem\\n# SZS output start CNFRefutation\\nfof(predefinitionsA15, axiom, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~(p__d__instance(X6,X4))|~(p__d__instance(X6,X5)))), file('adimen.sumo.tstp', predefinitionsA15)).\\nfof(predefinitionsA24, axiom, ((((![X7, X8, X9]:(p__d__disjointDecomposition3(X7,X8,X9)<=>p__d__disjoint(X8,X9))&![X7, X8, X9, X10]:(p__d__disjointDecomposition4(X7,X8,X9,X10)<=>((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X9,X10))))&![X7, X8, X9, X10, X11]:(p__d__disjointDecomposition5(X7,X8,X9,X10,X11)<=>(((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X10,X11))))&![X7, X8, X9, X10, X11, X12]:(p__d__disjointDecomposition6(X7,X8,X9,X10,X11,X12)<=>(((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X11,X12))))&![X7, X8, X9, X10, X11, X12, X13]:(p__d__disjointDecomposition7(X7,X8,X9,X10,X11,X12,X13)<=>((((((((((((((p__d__disjoint(X8,X9)&p__d__disjoint(X8,X10))&p__d__disjoint(X8,X11))&p__d__disjoint(X8,X12))&p__d__disjoint(X8,X13))&p__d__disjoint(X9,X10))&p__d__disjoint(X9,X11))&p__d__disjoint(X9,X12))&p__d__disjoint(X9,X13))&p__d__disjoint(X10,X11))&p__d__disjoint(X10,X12))&p__d__disjoint(X10,X13))&p__d__disjoint(X11,X12))&p__d__disjoint(X11,X13))&p__d__disjoint(X12,X13)))), file('adimen.sumo.tstp', predefinitionsA24)).\\nfof(whiteBoxTruthTest0005, conjecture, ?[X7]:~(p__d__instance(X7,c__SetOrClass)), file('./goals/whitebox/whiteBoxTruthTest0005.adimen.tstp', whiteBoxTruthTest0005)).\\nfof(mergeA331, axiom, p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition), file('adimen.sumo.tstp', mergeA331)).\\nfof(miloA4203, axiom, p__d__instance(c__AmericansWithDisabilitiesAct,c__Proposition), file('adimen.sumo.tstp', miloA4203)).\\nfof(c_0_5, plain, ![X4, X5]:(p__d__disjoint(X4,X5)<=>![X6]:(~p__d__instance(X6,X4)|~p__d__instance(X6,X5))), inference(fof_simplification,[status(thm)],[predefinitionsA15])).\\nfof(c_0_6, plain, ![X7870, X7871, X7872, X7873, X7874, X7875, X7876, X7877, X7878, X7879, X7880, X7881, X7882, X7883, X7884, X7885, X7886, X7887, X7888, X7889, X7890, X7891, X7892, X7893, X7894]:((((((~p__d__disjointDecomposition3(X7870,X7871,X7872)|p__d__disjoint(X7871,X7872))&(~p__d__disjoint(X7871,X7872)|p__d__disjointDecomposition3(X7870,X7871,X7872)))&((((p__d__disjoint(X7874,X7875)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))&(p__d__disjoint(X7874,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(p__d__disjoint(X7875,X7876)|~p__d__disjointDecomposition4(X7873,X7874,X7875,X7876)))&(~p__d__disjoint(X7874,X7875)|~p__d__disjoint(X7874,X7876)|~p__d__disjoint(X7875,X7876)|p__d__disjointDecomposition4(X7873,X7874,X7875,X7876))))&(((((((p__d__disjoint(X7878,X7879)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))&(p__d__disjoint(X7878,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7878,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7880)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7879,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(p__d__disjoint(X7880,X7881)|~p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881)))&(~p__d__disjoint(X7878,X7879)|~p__d__disjoint(X7878,X7880)|~p__d__disjoint(X7878,X7881)|~p__d__disjoint(X7879,X7880)|~p__d__disjoint(X7879,X7881)|~p__d__disjoint(X7880,X7881)|p__d__disjointDecomposition5(X7877,X7878,X7879,X7880,X7881))))&(((((((((((p__d__disjoint(X7883,X7884)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))&(p__d__disjoint(X7883,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7883,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7885)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7884,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7886)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7885,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(p__d__disjoint(X7886,X7887)|~p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887)))&(~p__d__disjoint(X7883,X7884)|~p__d__disjoint(X7883,X7885)|~p__d__disjoint(X7883,X7886)|~p__d__disjoint(X7883,X7887)|~p__d__disjoint(X7884,X7885)|~p__d__disjoint(X7884,X7886)|~p__d__disjoint(X7884,X7887)|~p__d__disjoint(X7885,X7886)|~p__d__disjoint(X7885,X7887)|~p__d__disjoint(X7886,X7887)|p__d__disjointDecomposition6(X7882,X7883,X7884,X7885,X7886,X7887))))&((((((((((((((((p__d__disjoint(X7889,X7890)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894))&(p__d__disjoint(X7889,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7889,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7891)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7890,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7892)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7891,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7893)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7892,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(p__d__disjoint(X7893,X7894)|~p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))&(~p__d__disjoint(X7889,X7890)|~p__d__disjoint(X7889,X7891)|~p__d__disjoint(X7889,X7892)|~p__d__disjoint(X7889,X7893)|~p__d__disjoint(X7889,X7894)|~p__d__disjoint(X7890,X7891)|~p__d__disjoint(X7890,X7892)|~p__d__disjoint(X7890,X7893)|~p__d__disjoint(X7890,X7894)|~p__d__disjoint(X7891,X7892)|~p__d__disjoint(X7891,X7893)|~p__d__disjoint(X7891,X7894)|~p__d__disjoint(X7892,X7893)|~p__d__disjoint(X7892,X7894)|~p__d__disjoint(X7893,X7894)|p__d__disjointDecomposition7(X7888,X7889,X7890,X7891,X7892,X7893,X7894)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[predefinitionsA24])])])])).\\nfof(c_0_7, negated_conjecture, ~(?[X7]:~p__d__instance(X7,c__SetOrClass)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[whiteBoxTruthTest0005])])).\\nfof(c_0_8, plain, ![X7779, X7780, X7781, X7782, X7783]:((~p__d__disjoint(X7779,X7780)|(~p__d__instance(X7781,X7779)|~p__d__instance(X7781,X7780)))&((p__d__instance(esk1_2(X7782,X7783),X7782)|p__d__disjoint(X7782,X7783))&(p__d__instance(esk1_2(X7782,X7783),X7783)|p__d__disjoint(X7782,X7783)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])).\\ncnf(c_0_9, plain, (p__d__disjoint(X1,X2)|~p__d__disjointDecomposition6(X3,X4,X5,X1,X6,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).\\ncnf(c_0_10, plain, (p__d__disjointDecomposition6(c__Abstract,c__Quantity,c__Attribute,c__SetOrClass,c__Relation,c__Proposition)), inference(split_conjunct,[status(thm)],[mergeA331])).\\nfof(c_0_11, negated_conjecture, ![X14607]:p__d__instance(X14607,c__SetOrClass), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).\\ncnf(c_0_12, plain, (~p__d__disjoint(X1,X2)|~p__d__instance(X3,X1)|~p__d__instance(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).\\ncnf(c_0_13, plain, (p__d__disjoint(c__SetOrClass,c__Proposition)), inference(spm,[status(thm)],[c_0_9, c_0_10])).\\ncnf(c_0_14, negated_conjecture, (p__d__instance(X1,c__SetOrClass)), inference(split_conjunct,[status(thm)],[c_0_11])).\\ncnf(c_0_15, plain, (p__d__instance(c__AmericansWithDisabilitiesAct,c__Proposition)), inference(split_conjunct,[status(thm)],[miloA4203])).\\ncnf(c_0_16, plain, (~p__d__instance(X1,c__Proposition)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])])).\\ncnf(c_0_17, plain, ($false), inference(sr,[status(thm)],[c_0_15, c_0_16]), ['proof']).\\n# SZS output end CNFRefutation\\n# Training examples: 0 positive, 0 negative\\n\""
  ]
}